High-altitude hypoxia exposure inhibits erythrophagocytosis by inducing macrophage ferroptosis in the spleen

Author Response

The following is the authors’ response to the original reviews.

We gratefully thank the editors and all reviewers for their time spend making their constructive remarks and useful suggestions, which has significantly raised the quality of the manuscript and has enable us to improve the manuscript. Each suggested comment brought forward by the reviewers was accurately considered. The manuscript has been revised in consideration of all suggestions.

Reviewer #1 (Public Review):

Wang and all present an interesting body of work focused on the effects of high altitude and hypoxia on erythropoiesis, resulting in erythrocytosis. This work is specifically focused on the spleen, identifying splenic macrophages as central cells in this effect. This is logical since these cells are involved in erythrophagocytosis and iron recycling. The results suggest that hypoxia induces splenomegaly with decreased number of splenic macrophages. There is also evidence that ferroptosis is induced in these macrophages, leading to cell destruction. Finally, the data suggest that ferroptosis in splenic red pulp macrophages causes the decrease in RBC clearance, resulting in erythrocytosis aka lengthening the RBC lifespan. However, there are many issues with the presented results, with somewhat superficial data, meaning the conclusions are overstated and there is decreased confidence that the hypotheses and observed results are directly causally related to hypoxia.

Major points:

1. The spleen is a relatively poorly understood organ but what is known about its role in erythropoiesis especially in mice is that it functions both to clear as well as to generate RBCs. The later process is termed extramedullary hematopoiesis and can occur in other bones beyond the pelvis, liver, and spleen. In mice, the spleen is the main organ of extramedullary erythropoiesis. The finding of transiently decreased spleen size prior to splenomegaly under hypoxic conditions is interesting but not well developed in the manuscript. This is a shortcoming as this is an opportunity to evaluate the immediate effect of hypoxia separately from its more chronic effect. Based just on spleen size, no conclusions can be drawn about what happens in the spleen in response to hypoxia.

Thank you for your insightful comments and questions. The spleen is instrumental in both immune response and the clearance of erythrocytes, as well as serving as a significant reservoir of blood in the body. This organ, characterized by its high perfusion rate and pliability, constricts under conditions of intense stress, such as during peak physical exertion, the diving reflex, or protracted periods of apnea. This contraction can trigger an immediate release of red blood cells (RBCs) into the bloodstream in instances of substantial blood loss or significant reduction of RBCs. Moreover, elevated oxygen consumption rates in certain animal species can be partially attributed to splenic contractions, which augment hematocrit levels and the overall volume of circulating blood, thereby enhancing venous return and oxygen delivery (Dane et al. J Appl Physiol, 2006, 101:289-97; Longhurst et al. Am J Physiol, 1986, 251: H502-9). In our investigation, we noted a significant contraction of the spleen following exposure to hypoxia for a period of one day. We hypothesized that the body, under such conditions, is incapable of generating sufficient RBCs promptly enough to facilitate enhanced oxygen delivery. Consequently, the spleen reacts by releasing its stored RBCs through splenic constriction, leading to a measurable reduction in spleen size.

However, we agree with you that further investigation is required to fully understand the implications of these changes. Considering the comments, we extended our research by incorporating more detailed examinations of spleen morphology and function during hypoxia, including the potential impact on extramedullary hematopoiesis. We anticipate that such an expanded analysis would not only help elucidate the initial response to hypoxia but also provide insights into the more chronic effects of this condition on spleen function and erythropoiesis.

1. Monocyte repopulation of tissue resident macrophages is a minor component of the process being described and it is surprising that monocytes in the bone marrow and spleen are also decreased. Can the authors conjecture why this is happening? Typically, the expectation would be that a decrease in tissue resident macrophages would be accompanied by an increase in monocyte migration into the organ in a compensatory manner.

We appreciate your insightful query regarding the observed decrease in monocytes in the bone marrow and spleen, particularly considering the typical compensatory increase in monocyte migration into organs following a decrease in tissue resident macrophages.

The observed decrease in monocytes within the bone marrow is likely attributable to the fact that monocytes and precursor cells for red blood cells (RBCs) both originate from the same hematopoietic stem cells within the bone marrow. It is well established that exposure to hypobaric hypoxia (HH) induces erythroid differentiation specifically within the bone marrow, originating from these hematopoietic stem cells (Exp Hematol, 2021 May;97:32-46). As such, the differentiation to monocyte is reduced under hypoxic conditions, which may subsequently cause a decrease in migration to spleen.

Furthermore, we hypothesize that an increased migration of monocytes to other tissues under HH exposure may also contribute to the decreased migration to the spleen. The liver, which partially contributes to the clearance of RBCs, may play a role in this process. Our investigations to date have indeed identified an increased monocyte migration to the liver. We were pleased to discover an elevation in CSF1 expression in the liver following HH exposure for both 7 and 14 days. This finding was corroborated through flow cytometry, which confirmed an increase in monocyte migration to the liver.

Consequently, we propose that under HH conditions, the liver requires an increased influx of monocytes, which in turn leads to a decrease in monocyte migration to the spleen. However, it is important to note that these findings will be discussed more comprehensively in our forthcoming publication, and as such, the data pertaining to these results have not been included in the current manuscript.

1. Figure 3 does not definitively provide evidence that cell death is specifically occurring in splenic macrophages and the fraction of Cd11b+ cells is not changed in NN vs HH. Furthermore, the IHC of F4/80 in Fig 3U is not definitive as cells can express F4/80 more or less brightly and no negative/positive controls are shown for this panel.

We appreciate your insightful comments and critiques regarding Figure 3. We acknowledge that the figure, as presented, does not definitively demonstrate that cell death is specifically occurring in splenic macrophages. While it is challenging to definitively determine the occurrence of cell death in macrophages based solely on Figure 3D-F, our single-cell analysis provides strong evidence that such an event occurs. We initially observed cell death within the spleen under hypobaric hypoxia (HH) conditions, and to discern the precise cell type involved, we conducted single-cell analyses. Regrettably, we did not articulate this clearly in our preliminary manuscript.

In the revised version, we have modified the sequence of Figure 3A-C and Figure 3D-F for better clarity. Besides, we observed a significant decrease in the fraction of F4/80hiCD11bhi macrophages under HH conditions compared to NN. To make the changes more evident in CD86 and CD206, we have transformed these scatter plots into histograms in our revised manuscript.

Considering the limitations of F4/80 as a conclusive macrophage identifier, we have concurrently presented the immunohistochemical (IHC) analyses of heme oxygenase-1 (HO-1). Functioning as a macrophage marker, particularly in cells involved in iron metabolism, HO-1 offers additional diagnostic accuracy. Observations from both F4/80 and HO-1 staining suggested a primary localization of positively stained cells within the splenic red pulp. Following exposure to hypoxia-hyperoxia (HH) conditions, a decrease was noted in the expression of both F4/80 and HO-1. This decrease implies that HH conditions contribute to a reduction in macrophage population and impede the iron metabolism process. In the revised version of our manuscript, we have enhanced the clarity of Figure 3U to illustrate the presence of positive staining, with an emphasis on HO-1 staining, which is predominantly observed in the red pulp.

1. The phagocytic function of splenic red pulp macrophages relative to infection cannot be used directly to understand erythrophagocytosis. The standard approach is to use opsonized RBCs in vitro. Furthermore, RBC survival is a standard method to assess erythrophagocytosis function. In this method, biotin is injected via tail vein directly and small blood samples are collected to measure the clearance of biotinilation by flow; kits are available to accomplish this. Because the method is standard, Fig 4D is not necessary and Fig 4E needs to be performed only in blood by sampling mice repeatedly and comparing the rate of biotin decline in HH with NN (not comparing 7 d with 14 d).

We appreciate your insightful comments and suggestions. We concur that the phagocytic function of splenic red pulp macrophages in the context of infection may not be directly translatable to understanding erythrophagocytosis. Given our assessment that the use of cy5.5-labeled E.coli alone may not be sufficient to accurately evaluate the phagocytic function of macrophages, we extended our study to include the use of NHS-biotin-labeled RBCs to assess phagocytic capabilities. While the presence of biotin-labeled RBCs in the blood could provide an indication of RBC clearance, this measure does not exclusively reflect the spleen's role in the process, as it fails to account for the clearance activities of other organs.

Consequently, we propose that the remaining biotin-labeled RBCs in the spleen may provide a more direct representation of the organ's function in RBC clearance and sequestration. Our observations of diminished erythrophagocytosis at both 7- and 14-days following exposure to HH guided our subsequent efforts to quantify biotin-labeled RBCs in both the circulatory system and spleen. These measurements were conducted during the 7 to 14-day span following the confirmation of impaired erythrophagocytosis. Comparative evaluation of RBC clearance rates under NN and HH conditions provided further evidence supporting our preliminary observations, with the data revealing a decrease in the RBC clearance rate in the context of HH conditions. In response to feedback from other reviewers, we have elected to exclude the phagocytic results and the diagram of the erythrocyte labeling assay. These amendments will be incorporated into the revised manuscript. The reviewers' constructive feedback has played a crucial role in refining the methodological precision and coherence of our investigation.

1. It is unclear whether Tuftsin has a specific effect on phagocytosis of RBCs without other potential confounding effects. Furthermore, quantifying iron in red pulp splenic macrophages requires alternative readily available more quantitative methods (e.g. sorted red pulp macrophages non-heme iron concentration).

We appreciate your comments and questions regarding the potential effect of Tuftsin on the phagocytosis of RBCs and the quantification of iron in red pulp splenic macrophages. Regarding the role of Tuftsin, we concur that the literature directly associating Tuftsin with erythrophagocytosis is scant. The work of Gino Roberto Corazza et al. does suggest a link between Tuftsin and general phagocytic capacity, but it does not specifically address erythrophagocytosis (Am J Gastroenterol, 1999;94:391-397). We agree that further investigations are required to elucidate the potential confounding effects and to ascertain whether Tuftsin has a specific impact on the phagocytosis of RBCs. Concerning the quantification of iron in red pulp splenic macrophages, we acknowledge your suggestion to employ readily available and more quantitative methods. We have incorporated additional Fe2+ staining in the spleen at two time points: 7 and 14 days subsequent to HH exposure (refer to the following Figure). The resultant data reveal an escalated deposition of Fe2+ within the red pulp, as evidenced in Figures 5 (panels L and M) and Figure S1 (panels L and M).

1. In Fig 5, PBMCs are not thought to represent splenic macrophages and although of some interest, does not contribute significantly to the conclusions regarding splenic macrophages at the heart of the current work. The data is also in the wrong direction, namely providing evidence that PBMCs are relatively iron poor which is not consistent with ferroptosis which would increase cellular iron.

We appreciate your insightful critique regarding Figure 5 and the interpretation of our data on peripheral blood mononuclear cells (PBMCs) in relation to splenic macrophages. We understand that PBMCs do not directly represent splenic macrophages, and we agree that any conclusions drawn from PBMCs must be considered with caution when discussing the behavior of splenic macrophages.

The primary rationale for incorporating PBMCs into our study was to investigate the potential correspondence between their gene expression changes and those observed in the spleen after HH exposure. This was posited as a working hypothesis for further exploration rather than a conclusive statement. The gene expression in PBMCs was congruous with changes in the spleen's gene expression, demonstrating an iron deficiency phenotype, ostensibly due to the mobilization of intracellular iron for hemoglobin synthesis. Thus, it is plausible that NCOA4 may facilitate iron mobilization through the degradation of ferritin to store iron.

It remains ambiguous whether ferroptosis was initiated in the PBMCs during our study. Ferroptosis primarily occurs as a response to an increase in Fe2+ rather than an overall increase in intracellular iron. Our preliminary proposition was that relative changes in gene expression in PBMCs could potentially mirror corresponding changes in protein expression in the spleen, thereby potentially indicating alterations in iron processing capacity post-HH exposure. However, we fully acknowledge that this is a conjecture requiring further empirical substantiation or clinical validation.

1. Tfr1 increase is typically correlated with cellular iron deficiency while ferroptosis consistent with iron loading. The direction of the changes in multiple elements relevant to iron trafficking is somewhat confusing and without additional evidence, there is little confidence that the authors have reached the correct conclusion. Furthermore, the results here are analyses of total spleen samples rather than specific cells in the spleen.

We appreciate your astute comments and agree that the observed increase in transferrin receptor (TfR) expression, typically associated with cellular iron deficiency, appears contradictory to the expected iron-loading state associated with ferroptosis. We understand that this apparent contradiction might engender some uncertainty about our conclusions. In our investigation, we evaluated total spleen samples as opposed to distinct cell types within the spleen, a factor that could have contributed to the seemingly discordant findings. An integral element to bear in mind is the existence of immature RBCs in the spleen, particularly within the hematopoietic island where these immature RBCs cluster around nurse macrophages. These immature RBCs contain abundant TfR which was needed for iron uptake and hemoglobin synthesis. These cells, which prove challenging to eliminate via perfusion, might have played a role in the observed upregulation in TfR expression, especially in the aftermath of HH exposure. Our further research revealed that the expression of TfR in macrophages diminished following hypoxic conditions, thereby suggesting that the elevated TfR expression in tissue samples may predominantly originate from other cell types, especially immature RBCs (refer to Author response image 5).

Reviewer #2 (Public Review):

The authors aimed at elucidating the development of high altitude polycythemia which affects mice and men staying in the hypoxic atmosphere at high altitude (hypobaric hypoxia; HH). HH causes increased erythropoietin production which stimulates the production of red blood cells. The authors hypothesize that increased production is only partially responsible for exaggerated red blood cell production, i.e. polycythemia, but that decreased erythrophagocytosis in the spleen contributes to high red blood cells counts.

The main strength of the study is the use of a mouse model exposed to HH in a hypobaric chamber. However, not all of the reported results are convincing due to some smaller effects which one may doubt to result in the overall increase in red blood cells as claimed by the authors. Moreover, direct proof for reduced erythrophagocytosis is compromised due to a strong spontaneous loss of labelled red blood cells, although effects of labelled E. coli phagocytosis are shown. Their discussion addresses some of the unexpected results, such as the reduced expression of HO-1 under hypoxia but due to the above-mentioned limitations much of the discussion remains hypothetical.

Thank you for your valuable feedback and insight. We appreciate the recognition of the strength of our study model, the exposure of mice to hypobaric hypoxia (HH) in a hypobaric animal chamber. We also understand your concerns about the smaller effects and their potential impact on the overall increase in red blood cells (RBCs), as well as the apparent reduced erythrophagocytosis due to the loss of labelled RBCs.

Erythropoiesis has been predominantly attributed to the amplified production of RBCs under conditions of HH. The focus of our research was to underscore the potential acceleration of hypoxia-associated polycythemia (HAPC) as a result of compromised erythrophagocytosis. Considering the spontaneous loss of labelled RBCs in vivo, we assessed the clearance rate of RBCs at the stages of 7 and 14 days within the HH environment, and subsequently compared this rate within the period from 7 to 14 days following the clear manifestation of erythrophagocytosis impairment at the two aforementioned points identified in our study. This approach was designed to negate the effects of spontaneous loss of labelled RBCs in both NN and HH conditions. Correspondingly, the results derived from blood and spleen analyses corroborated a decline in the RBC clearance rate under HH when juxtaposed with NN conditions.

Apart from the E. coli phagocytosis and the labeled RBCs experiment (this part of the results was removed in the revision), the injection of Tuftsin further substantiated the impairment of erythrophagocytosis in the HH spleen, as evidenced by the observed decrease in iron within the red pulp of the spleen post-perfusion. Furthermore, to validate our findings, we incorporated RBCs staining in splenic cells at 7 and 14 days of HH exposure, which provided concrete confirmation of impaired erythrophagocytosis (new Figure 4E).

As for the reduced expression of heme oxygenase-1 (HO-1) under hypoxia, we agree that this was an unexpected result, and we are in the process of further exploring the underlying mechanisms. It is possible that there are other regulatory pathways at play that are yet to be identified. However, we believe that by offering possible interpretations of our data and potential directions for future research, we contribute to the ongoing scientific discourse in this area.

Reviewer #3 (Public Review):

The manuscript by Yang et al. investigated in mice how hypobaric hypoxia can modify the RBC clearance function of the spleen, a concept that is of interest. Via interpretation of their data, the authors proposed a model that hypoxia causes an increase in cellular iron levels, possibly in RPMs, leading to ferroptosis, and downregulates their erythrophagocytic capacity. However, most of the data is generated on total splenocytes/total spleen, and the conclusions are not always supported by the presented data. The model of the authors could be questioned by the paper by Youssef et al. (which the authors cite, but in an unclear context) that the ferroptosis in RPMs could be mediated by augmented erythrophagocytosis. As such, the loss of RPMs in vivo which is indeed clear in the histological section shown (and is a strong and interesting finding) can be not directly caused by hypoxia, but by enhanced RBC clearance. Such a possibility should be taken into account.

Thank you for your insightful comments and constructive feedback. In their research, Youssef et al. (2018) discerned that elevated erythrophagocytosis of stressed red blood cells (RBCs) instigates ferroptosis in red pulp macrophages (RPMs) within the spleen, as evidenced in a mouse model of transfusion. This augmentation of erythrophagocytosis was conspicuous five hours post-injection of RBCs. Conversely, our study elucidated the decrease in erythrophagocytosis in the spleen after both 7 and 14 days.

Typically, macrophages exhibit an enhanced phagocytic capacity in the immediate aftermath of stress or stimulation. Nonetheless, the temporal points of observation in our study were considerably extended (7 and 14 days). It is currently unclear whether the phagocytic capacity is amplified during the acute phase of HH exposure, especially on the first day. Considering that the spleen contraction on the next day of HH leads to the release of stored RBCs into the bloodstream, and whether this initial reaction leads to ferroptosis, and the phagocytic capacity of RBCs is subsequently weakened after 7 or 14 days under sustained HH conditions.

Major points:

1. The authors present data from total splenocytes and then relate the obtained data to RPMs, which are quantitatively a minor population in the spleen. Eg, labile iron is increased in the splenocytes upon HH, but the manuscript does not show that this occurs in the red pulp or RPMs. They also measure gene/protein expression changes in the total spleen and connect them to changes in macrophages, as indicated in the model Figure (Fig. 7). HO-1 and levels of Ferritin (L and H) can be attributed to the drop in RPMs in the spleen. Are any of these changes preserved cell-intrinsically in cultured macrophages? This should be shown to support the model (relates also to lines 487-88, where the authors again speculate that hypoxia decreases HO-1 which was not demonstrated). In the current stage, for example, we do not know if the labile iron increase in cultured cells and in the spleen in vivo upon hypoxia is the same phenomenon, and why labile iron is increased. To improve the manuscript, the authors should study specifically RPMs.

We express our gratitude for your perceptive remarks. In our initial manuscript, we did not evaluate labile iron within the red pulp and red pulp macrophages (RPMs). To address this oversight, we utilized the Lillie staining method, in accordance with the protocol outlined by Liu et al., (Chemosphere, 2021, 264(Pt 1):128413), to discern Fe2+ presence within these regions. The outcomes were consistent with our antecedent Western blot and flow cytometry findings in the spleen, corroborating an increment in labile iron specifically within the red pulp of the spleen.

However, we acknowledge the necessity for other supplementary experimental efforts to further validate these findings. Additionally, we scrutinized the expression of heme oxygenase-1 (HO-1) and iron-related proteins, including transferrin receptor (TfR), ferroportin (Fpn), ferritin (Ft), and nuclear receptor coactivator 4 (NCOA4) in primary macrophages subjected to 1% hypoxic conditions, both with and without hemoglobin treatment. Our results indicated that the expression of ferroptosis-related proteins was consistent with in vivo studies, however the expression of iron related proteins was not similar in vitro and in vivo. It suggesting that the increase in labile iron in cultured cells and the spleen in vivo upon hypoxia are not identical phenomena. However, the precise mechanism remains elusive.

In our study, we observed a decrease in HO-1 protein expression following 7 and 14 days of HH exposure, as shown in Figure 3U, 5A, and S1A. This finding contradicts previous research that identified HO-1 as a hypoxia-inducible factor (HIF) target under hypoxic conditions (P J Lee et al., 1997). Our discussion, therefore, addressed the potential discrepancy in HO-1 expression under HH. According to our findings, HO-1 regulation under HH appears to be predominantly influenced by macrophage numbers and the RBCs to be processed in the spleen or macrophages, rather than by hypoxia alone.

It is challenging to discern whether the increased labile iron observed in vitro accurately reflects the in vivo phenomenon, as replicating the iron requirements for RBCs production induced by HH in vitro is inherently difficult. However, by integrating our in vivo and in vitro studies, we determined that the elevated Fe2+ levels were not dependent on HO-1 protein expression, as HO-1 levels was increased in vitro while decreasing in vivo under hypoxic/HH exposure.

1. The paper uses flow cytometry, but how this method was applied is suboptimal: there are no gating strategies, no indication if single events were determined, and how cell viability was assessed, which are the parent populations when % of cells is shown on the graphs. How RBCs in the spleen could be analyzed without dedicated cell surface markers? A drop in splenic RPMs is presented as the key finding of the manuscript but Fig. 3M shows gating (suboptimal) for monocytes, not RPMs. RPMs are typically F4/80-high, CD11-low (again no gating strategy is shown for RPMs). Also, the authors used single-cell RNAseq to detect a drop in splenic macrophages upon HH, but they do not indicate in Fig. A-C which cluster of cells relates to macrophages. Cell clusters are not identified in these panels, hence the data is not interpretable).

Thank you for your comments and constructive critique regarding our flow cytometry methodology and presentation. We understand the need for greater transparency and detailed explanation of our procedures, and we acknowledge that the lack of gating strategies and other pertinent information in our initial manuscript may have affected the clarity of our findings.

In our initial report, we provided an overview of the decline in migrated macrophages (F4/80hiCD11bhi), including both M1 and M2 expression in migrated macrophages, as illustrated in Figure 3, but did not specifically address the changes in red pulp macrophages (RPMs). Based on previous results, it is difficult to identify CD11b- and CD11blo cells. We will repeat the results and attempt to identify F4/80hiCD11blo cells in the revised manuscript. The results of the reanalysis are now included (Figure 3M). However, single-cell in vivo analysis studies may more accurately identify specific cell types that decrease after exposure to HH.

Furthermore, we substantiated the reduction in red pulp, as evidenced by Figure 4J, given that iron processing primarily occurs within the red pulp. In Figure 3, our initial objective was merely to illustrate the reduction in total macrophages in the spleen following HH exposure.

To further clarify the characterization of various cell types, we conducted a single-cell analysis. Our findings indicated that clusters 0,1,3,4,14,18, and 29 represented B cells, clusters 2, 10, 12, and 28 represented T cells, clusters 15 and 22 corresponded to NK cells, clusters 5, 11, 13, and 19 represented NKT cells, clusters 6, 9, and 24 represented cell cycle cells, clusters 26 and 17 represented plasma cells, clusters 21 and 23 represented neutrophils, cluster 30 represented erythrocytes, and clusters 7, 8, 16, 20, 24, and 27 represented dendritic cells (DCs) and macrophages, as depicted in Figure 3E.

1. The authors draw conclusions that are not supported by the data, some examples: a) they cannot exclude eg the compensatory involvement of the liver in the RBCs clearance (the differences between HH sham and HH splenectomy is mild in Fig. 2 E, F and G).

Thank you for your insightful comments and for pointing out the potential involvement of other organs, such as the liver, in the RBC clearance under HH conditions. We concur with your observation that the differences between the HH sham and HH splenectomy conditions in Fig. 2 E, F, and G are modest. This could indeed suggest a compensatory role of other organs in RBC clearance when splenectomy is performed. Our intent, however, was to underscore the primary role of the spleen in this process under HH exposure.

In fact, after our initial investigations, we conducted a more extensive study examining the role of the liver in RBC clearance under HH conditions. Our findings, as illustrated in the figures submitted with this response, indeed support a compensatory role for the liver. Specifically, we observed an increase in macrophage numbers and phagocytic activity in the liver under HH conditions. Although the differences in RBC count between the HH sham and HH splenectomy conditions may seem minor, it is essential to consider the unit of this measurement, which is value*1012/ml. Even a small numerical difference can represent a significant biological variation at this scale.

b) splenomegaly is typically caused by increased extramedullary erythropoiesis, not RBC retention. Why do the authors support the second possibility? Related to this, why do the authors conclude that data in Fig. 4 G,H support the model of RBC retention? A significant drop in splenic RBCs (poorly gated) was observed at 7 days, between NN and HH groups, which could actually indicate increased RBC clearance capacity = less retention.

Prior investigations have predominantly suggested that spleen enlargement under hypoxic conditions stems from the spleen's extramedullary hematopoiesis. Nevertheless, an intriguing study conducted in 1994 by the General Hospital of Xizang Military Region reported substantial exaggeration and congestion of splenic sinuses in high altitude polycythemia (HAPC) patients. This finding was based on the dissection of spleens from 12 patients with HAPC (Zou Xunda, et al., Southwest Defense Medicine, 1994;5:294-296). Moreover, a recent study indicated that extramedullary erythropoiesis reaches its zenith between 3 to 7 days (Wang H et al., 2021).

Considering these findings, the present study postulates that hypoxia-induced inhibition of erythrophagocytosis may lead to RBC retention. However, we acknowledge that the manuscript in its current preprint form does not offer conclusive evidence to substantiate this hypothesis. To bridge this gap, we further conducted experiments where the spleen was perfused, and total cells were collected post HH exposure. These cells were then smeared onto slides and subjected to Wright staining. Our results unequivocally demonstrate an evident increase in deformation and retention of RBCs in the spleen following 7 and 14 days of HH exposure. This finding strengthens our initial hypothesis and contributes a novel perspective to the understanding of splenic responses under hypoxic conditions.

c) lines 452-54: there is no data for decreased phagocytosis in vivo, especially in the context of erythrophagocytosis. This should be done with stressed RBCs transfusion assays, very good examples, like from Youssef et al. or Threul et al. are available in the literature.

Thanks. In their seminal work, Youssef and colleagues demonstrated that the transfusion of stressed RBCs triggers erythrophagocytosis and subsequently incites ferroptosis in red pulp macrophages (RPMs) within a span of five hours. Given these observations, the applicability of this model to evaluate macrophage phagocytosis in the spleen or RPMs under HH conditions may be limited, as HH has already induced erythropoiesis in vivo. In addition, it was unclear whether the membrane characteristics of stress induced RBCs were similar to those of HH induced RBCs, as this is an important signal for in vivo phagocytosis. The ambiguity arises from the fact that we currently lack sufficient knowledge to discern whether the changes in phagocytosis are instigated by the presence of stressed RBCs or by changes of macrophages induced by HH in vivo. Nonetheless, we appreciate the potential value of this approach and intend to explore its utility in our future investigations. The prospect of distinguishing the effects of stressed RBCs from those of HH on macrophage phagocytosis is an intriguing line of inquiry that could yield significant insights into the mechanisms governing these physiological processes. We will investigate this issue in our further study.

d) Line 475 - ferritinophagy was not shown in response to hypoxia by the manuscript, especially that NCOA4 is decreased, at least in the total spleen.

Drawing on the research published in eLife in 2015, it was unequivocally established that ferritinophagy, facilitated by Nuclear Receptor Coactivator 4 (NCOA4), is indispensable for erythropoiesis. This process is modulated by iron-dependent HECT and RLD domain containing E3 ubiquitin protein ligase 2 (HERC2)-mediated proteolysis (Joseph D Mancias et al., eLife. 2015; 4: e10308). As is widely recognized, NCOA4 plays a critical role in directing ferritin (Ft) to the lysosome, where both NCOA4 and Ft undergo coordinated degradation. In our study, we provide evidence that exposure to HH stimulates erythropoiesis (Figure 1). We propose that this, in turn, could promote ferritinophagy via NCOA4, resulting in a decrease in NCOA4 protein levels post-HH exposure. We will further increase experiments to verify this concern. This finding not only aligns with the established understanding of ferritinophagy and erythropoiesis but also adds a novel dimension to the understanding of cellular responses to hypoxic conditions.

1. In a few cases, the authors show only representative dot plots or histograms, without quantification for n>1. In Fig. 4B the authors write about a significant decrease (although with n=1 no statistics could be applied here; of note, it is not clear what kind of samples were analyzed here). Another example is Fig. 6I. In this case, it is even more important as the data are conflicting the cited article and the new one: PMCID: PMC9908853 which shows that hypoxia stimulates efferocytosis. Sometimes the manuscript claim that some changes are observed, although they are not visible in representative figures (eg for M1 and M2 macrophages in Fig. 3M)

We recognize that our initial portrayal of Figure 4B was lacking in precision, given that it did not include the corresponding statistical graph. While our results demonstrated a significant reduction in the ability to phagocytose E. coli, in line with the recommendations of other reviewers, we have opted to remove the results pertaining to E. coli phagocytosis in this revision, as they primarily reflected immune function.

In relation to PMC9908853, which reported metabolic adaptation facilitating enhanced macrophage efferocytosis in limited-oxygen environments, it is worth noting that the macrophages investigated in this study were derived from ER-Hoxb8 macrophage progenitors following the removal of β-estradiol. Consequently, questions arise regarding the comparability between these cultured macrophages and primary macrophages obtained fresh from the spleen post HH exposure. The characteristics and functions of these two different macrophage sources may not align precisely, and this distinction necessitates further investigation.

1. There are several unclear issues in methodology:

what is the purity of primary RPMs in the culture? RPMs are quantitatively poorly represented in splenocyte single-cell suspensions. This reviewer is quite skeptical that the processing of splenocytes from approx 1 mm3 of tissue was sufficient to establish primary RPM cultures. The authors should prove that the cultured cells were indeed RPMs, not monocyte-derived macrophages or other splenic macrophage subtypes.

Thank you for your thoughtful comments and inquiries. Firstly, I apologize if we did not make it clear in the original manuscript. The purity of the primary RPMs in our culture was found to be approximately 40%, as identified by F4/80hiCD11blo markers using flow cytometry. We recognize that RPMs are typically underrepresented in splenocyte single-cell suspensions, and the concern you raise about the potential for contamination by other cell types is valid.

We apologize for any ambiguities in the methodological description that may have led to misunderstandings during the review. Indeed, the entirety of the spleen is typically employed for splenic macrophage culture. The size of the spleen can vary dependent on the species and age of the animal, but in mice, it is commonly approximately 1 cm in length. The spleen is then dissected into minuscule fragments, each approximately 1 mm3 in volume, to aid in enzymatic digestion. This procedure does not merely utilize a single 1 mm3 tissue fragment for RPMs cultures. Although the isolation and culture of spleen macrophages can present considerable challenges, our method has been optimized to enhance the yield of this specific cell population.

(around line 183) In the description of flow cytometry, there are several missing issues. In 1) it is unclear which type of samples were analyzed. In 2) it is not clear how splenocyte cell suspension was prepared.

1. Whole blood was extracted from the mice and collected into an anticoagulant tube, which was then set aside for subsequent thiazole orange (TO) staining.

2. Splenic tissue was procured from the mice and subsequently processed into a single-cell suspension using a 40 μm filter. The erythrocytes within the entire sample were subsequently lysed and eliminated, and the remaining cell suspension was resuspended in phosphate-buffered saline (PBS) in preparation for ensuing analyses.

We have meticulously revised these methodological details in the corresponding section of the manuscript to ensure clarity and precision.

In line 192: what does it mean: 'This step can be omitted from cell samples'?

The methodology employed for the quantification of intracellular divalent iron content and lipid peroxidation level was executed as follows: Splenic tissue was first processed into a single cell suspension, subsequently followed by the lysis of RBCs. It should be noted that this particular stage is superfluous when dealing with isolated cell samples. Subsequently, a total of 1 × 106 cells were incubated with 100 μL of BioTracker Far-red Labile Fe2+ Dye (1 mM, Sigma, SCT037, USA) for a duration of 1 hour, or alternatively, C11-Bodipy 581/591 (10 μM, Thermo Fisher, D3861, USA) for a span of 30 minutes. Post incubation, cells were thoroughly washed twice with PBS. Flow cytometric analysis was subsequently performed, utilizing the FL6 (638 nm/660 nm) channel for the determination of intracellular divalent iron content, and the FL1 (488 nm/525 nm) channel for the quantification of the lipid peroxidation level.

'TO method' is not commonly used anymore and hence it was unclear to this Reviewer. Reticulocytes should be analyzed with proper gating, using cell surface markers.

We are appreciative of your astute observation pertaining to the methodology we employed to analyze reticulocytes in our study. We value your recommendation to utilize cell surface markers for effective gating, which indeed represents a more modern and accurate approach. However, as reticulocyte identification is not the central focus of our investigation, we opted for the TO staining method—due to its simplicity and credibility of results. In our initial exploration, we adopted the TO staining method in accordance with the protocol outlined (Sci Rep, 2018, 8(1):12793), primarily owing to its established use and demonstrated efficacy in reticulocyte identification.

The description of 'phagocytosis of E. coli and RBCs' in the Methods section is unclear and incomplete. The Results section suggests that for the biotinylated RBCs, phagocytosis? or retention? Of RBCs was quantified in vivo, upon transfusion. However, the Methods section suggests either in vitro/ex vivo approach. It is vague what was indeed performed and how in detail. If RBC transfusion was done, this should be properly described. Of note, biotinylation of RBCs is typically done in vivo only, being a first step in RBC lifespan assay. The such assay is missing in the manuscript. Also, it is not clear if the detection of biotinylated RBCs was performed in permeablized cells (this would be required).

Thanks for the comments. In our initial methodology, we employed Cy5.5-labeled Escherichia coli to probe phagocytic function, albeit with the understanding that this may not constitute the most ideal model for phagocytosis detection within this context (in light of recommendations from other reviewers, we have removed the E. coli phagocytosis results from this revision, as they predominantly mirror immune function). Our fundamental aim was to ascertain whether HH compromises the erythrophagocytic potential of splenic macrophages. In pursuit of this, we subsequently analyzed the clearance of biotinylated RBCs in both the bloodstream and spleen to assess phagocytic functionality in vivo.

In the present study, instead of transfusing biotinylated RBCs into mice, we opted to inject N-Hydroxysuccinimide (NHS)-biotin into the bloodstream. NHS-biotin is capable of binding with cell membranes in vivo and can be recognized by streptavidin-fluorescein isothiocyanate (FITC) after cells are extracted from the blood or spleen in vitro. Consequently, biotin-labeled RBCs were detectable in both the blood and spleen following NHS-biotin injection for a duration of 21 days. Ultimately, we employed flow cytometry to analyze the NHS-biotin labeled RBCs in the blood or spleen. This method facilitates the detection of live cells and is not applicable to permeabilized cells. We believe this approach better aligns with our investigative goals and offers a more robust evaluation of erythrophagocytic function under hypoxic conditions.

Recommendations for the authors: please note that you control which, if any, revisions, to undertake.

Thank you for your comments and recommendations. We appreciate your understanding that the choice of implementing revisions ultimately rests with us. However, we also value your expertise and will seriously consider your suggestions as they can provide additional perspectives to our work and contribute to the overall quality and robustness of our study.

We strive to produce research that meets the highest scientific standards and we believe that constructive criticism, such as yours, helps us to achieve this objective. We will carefully review your comments and consider the appropriate changes to make in order to address your concerns and improve our manuscript.

Reviewer #1 (Recommendations For The Authors):

Minor:

1. HCV in text is a typo, should be HCT. Please edit.

Thanks for the correction. We’ve revised it.

1. Fig 2D is not useful beyond the more accurate measure of HCT in Fig 2G and should be removed.

Thank you for your feedback and suggestion about Fig. 2D. We understand your point regarding the comparative accuracy of HCT in Fig. 2G. However, our intention in including Fig. 2D was to provide a more intuitive visual representation of the erythrocyte position levels, which we believe complements the more precise HCT data. We have observed that the erythrocyte positions significantly increased for 14 days after HH splenectomy, and this trend is visually depicted in Fig. 2D. While HCT provides a more accurate measure, Fig. 2D provides a snapshot that can be more immediately graspable, especially for readers who may prefer visual data. Nevertheless, we appreciate your perspective and will reassess whether the inclusion of Fig. 2D adds enough value to the overall understanding of our findings. If we find that it indeed does not contribute significantly, we will consider removing it in line with your suggestion.

1. What is the purpose of performing splenectomy? It is well established that reticuloendothelial cells of the liver perform a redundant function to splenic macrophages and since these cells are not being evaluated, data following splenectomy is of limited value. Please remove or move to supplement. Alternatively, evaluate what happens in the liver in response to hypoxia. Is there an increase in erythroblasts? Is there a decrease in liver macrophages in the same way as in the spleen in non-splenectomized mice? The minimally increased HCT in hypoxic splenectomized mice (relative to non-splenectomized mice) suggests that the spleen does the primary work of clearance but not exclusively since there is still a major increase in response to hypoxia in splenectomized mice. The sentence (page 16, line 292) states that the spleen is essential which is not the case based on this data.

Thank you for your comments and recommendations. In reality, we have been consistently studying the liver's response to hypobaric hypoxia (HH) exposure. Nevertheless, the changes observed in the liver are contrary to those in the spleen, including an increase in macrophage count and the capacity for erythrophagocytosis, as well as processing heme iron (refer to the above figure for details).

It is widely accepted that HH exposure predominantly induces erythropoiesis by stimulating bone marrow production. The primary objective of this study was not to refute this central mechanism behind erythrocytosis. Instead, our intent was to supplement this understanding by proposing that impaired clearance of red blood cells (RBCs) could potentially exacerbate erythrocytosis. We believe this additional perspective could significantly enhance our understanding of the complex dynamics involved in RBC production and clearance under hypoxic conditions.

Reviewer #2 (Recommendations For The Authors):

The following questions and remarks should be considered by the authors:

1). The methods should clearly state whether the HH was discontinued during the 7- or 14-day exposure for cleaning, fresh water etc. Moreover, how was CO2 controlled? The procedure for splenectomy needs to be described in the methods.

Thank you for your insightful comments and questions. We apologize for any lack of clarity in our original description. To address your questions:

During the 7- or 14-day HH exposure, the HH was not discontinued for cleaning or providing fresh water. We ensured that the cage was thoroughly cleaned, and food and water were sufficiently stocked before placing the mice into the HH chamber. The design of the cage and the HH chamber allowed the mice to have continuous access to food and water during the entire exposure period.

Regarding the control of CO2, the HH chamber was equipped with a CO2 scrubbing system. The system utilized soda lime to absorb excess CO2 produced by the mice, and the air inside the chamber was exchanged with the air outside 25 times per hour to maintain a stable atmospheric concentration and ensure adequate oxygen supply.

As for the procedure for splenectomy, we apologize for the omission in the original manuscript. The mice were anesthetized using isoflurane, and a small incision was made in the left flank to expose the spleen. The spleen was then gently exteriorized, ligated, and excised. The incision was sutured, and the mice were allowed to recover under close monitoring. We ensured that all procedures were performed in accordance with our institution's guidelines for animal care.

1. The lack of changes in MCH needs explanation? During stress erythropoiesis some limit in iron availability should cause MCH decrease particularly if the authors claim that macrophages for rapid iron recycling are decreased. Fig 1A is dispensable. Fig 1G NN control 14 days does not make sense since it is higher than 7 days of HH.

Thank you for your insightful comments and queries. Regarding the lack of changes in Mean Corpuscular Hemoglobin (MCH), our hypothesis is that the decrease in iron recycling in the spleen following HH is potentially compensated by the increased iron absorption or supply from the liver, thus maintaining the iron requirement for erythropoiesis. This may explain why MCH levels did not significantly change after HH exposure. We have indeed observed an increase in macrophage numbers and their erythrophagocytosis/heme iron processing ability after HH exposure for 7 or 14 days in liver (please refer to the above figure for details), suggesting a compensatory mechanism to ensure adequate iron for erythropoiesis.

Regarding your comment on Fig 1A, we included this figure to provide a baseline of the experimental condition before any treatment. However, we understand your point and will consider removing it if it does not contribute significantly to the interpretation of our results. As for Fig 1G, we agree that the control at 14 days being higher than 7 days of HH may seem counterintuitive. We believe this could be due to individual variations among the mice or potential experimental errors. However, considering recommendations from other reviewers, we have removed this result from the revised manuscript.

1. Fig 2, the difference between sham and splenectomy is really marginal and not convincing. Is there also a difference at 7 days? Why does the spleen size decrease between 7 and 14 days?

We understand your concerns regarding the observed differences in Fig. 2 between sham and splenectomy groups. We acknowledge that while the absolute numerical differences may appear marginal, it is important to consider the unit of measurement. In the case of RBC count, the unit is 1012/L, hence even slight numerical differences can translate to significant variations in the actual count of RBCs.

We did not examine alterations occurring 7 days post-splenectomy in our study. The discernible trend of spleen size diminution between the 7th and 14th days is indeed compelling. It is plausible that this might be attributable to the body's adaptive response to hypobaric hypoxia (HH) exposure, wherein spleen size initially enlarges (at day 7) in response to compensatory erythropoiesis, followed by a reduction (at day 14) as the body acclimatizes to the HH conditions. Nevertheless, we did not identify a statistically significant difference between the measurements at day 7 and day 14, suggesting that this observation warrants further scrutiny.

1. Fig 3B, the clusters should be explained in detail. If the decrease in macrophages in Fig 3K/L is responsible for the effect, why does splenectomy not have a much stronger effect? How do the authors know which cells died in the calcein stained population in Fig 3D?

Thank you for your insightful queries and comments. Regarding Fig. 3B, we apologize for not providing sufficient detail on the clusters in the original manuscript. We will ensure that we include a comprehensive explanation of the clusters, including the specific cell types and their respective markers, in our revision. (clusters 0,1,3,4,14,18, and 29 represented B cells, clusters 2, 10, 12, and 28 represented T cells, clusters 15 and 22 corresponded to NK cells, clusters 5, 11, 13, and 19 represented NKT cells, clusters 6, 9, and 24 represented cell cycle cells, clusters 26 and 17 represented plasma cells, clusters 21 and 23 represented neutrophils, cluster 30 represented erythrocytes, and clusters 7, 8, 16, 20, 24, and 27 represented dendritic cells (DCs) and macrophages).

As for the decrease in macrophages observed in Fig. 3K/L, it's important to note that the spleen is a complex organ comprising numerous cell types, all of which can contribute to its overall function. While macrophages play a crucial role in iron recycling and erythropoiesis, other cell types and factors may also influence these processes. Therefore, while splenectomy results in the removal of all splenic cells, the overall impact on these processes may not be as pronounced as the specific reduction in macrophages due to compensatory mechanisms from other tissues and cells.

Concerning Fig. 3D, we acknowledge the ambiguity in the initial interpretation. The calcein staining was utilized to determine cell viability, but it doesn't identify the specific cell types that have died. To address this, we performed a single-cell analysis, which can provide a more accurate identification of the specific cell types affected.

1. Is the reduced phagocytic capacity in Fig4B significant? Erythrophagocytosis is compromised due to the considerable spontaneous loss of labelled erythrocytes; could other assays help? (potentially by a modified Chromium release assay?). Is it necessary to stimulated phagocytosis to see a significant effect?

We express our gratitude for your insightful queries and recommendations. In response to your initial question, the observed reduction in phagocytic capacity illustrated in Fig. 4B was indeed statistically significant. However, in alignment with feedback from other reviewers, we have elected to exclude the phagocytic results from this revised manuscript, as they predominantly reflect immune function rather than erythrophagocytosis of macrophages.

With respect to your proposal of potential alternatives to the erythrophagocytosis assay, we concur that the spontaneous loss of labeled erythrocytes could have influenced our results. Your suggestion of implementing a modified Chromium release assay is indeed an intriguing possibility that warrants further exploration.

Regarding the requirement for stimulating phagocytosis, we employed stimulation as a mechanism to investigate the potential for augmenting erythrophagocytosis and iron processing within the red pulp. Our findings suggest that increased phagocytosis in the spleen contributes positively to these processes. As part of the Tuftsin injection experiment, we assessed the RBC count and hemoglobin content. Despite an observed reduction trend, there were no statistically significant alterations. We are uncertain if the observation period was insufficiently long. Nevertheless, we concur that it would be worthwhile to explore inherent changes without external stimulation, and we will take this into consideration in our future research.

1. Can the observed ferroptosis be influenced by bi- and not trivalent iron chelators?

Thank you for your insightful question. Indeed, the role of iron chelators in the observed ferroptosis is an important aspect to explore. Ferroptosis is a form of regulated cell death characterized by an iron-dependent accumulation of lipid peroxides, and the role of different iron chelators could potentially influence this process.

In the case of bi- versus trivalent iron chelators, their influence on ferroptosis could be distinct due to their specificities for different forms of iron. However, we have not yet investigated this in our current study.

Your suggestion has highlighted a valuable direction for our future research. We agree that examining the influence of bi- and trivalent iron chelators on the observed ferroptosis would provide a deeper understanding of the iron-dependent mechanisms involved in this process. We will consider this important aspect in our subsequent investigations.

Reviewer #3 (Recommendations For The Authors):

Methodology:

1. Several syntax and grammatical errors, and unclear phrasing. Some factual errors as well: eg, line 380-81 the authors wrote that hypoxia increased viable cell numbers and phagocytosis ability, although their data suggest the opposite. Lines in Discussion 454-55 and in the Results 346-47 convey opposite messages.

We appreciate your attention to detail and your feedback on the language and factual discrepancies within the manuscript.

Upon revisiting lines 380-381, we would like to clarify that we had made a mistake. Our data indeed suggest that hypoxia led to a reduction in viable cell numbers and phagocytosis ability, not an increase as originally stated. We sincerely apologize for the confusion and will correct this statement in our revised manuscript.

As for the opposing messages between lines 454-455 in the Discussion and 346-347 in the Results, we apologize for any confusion caused. We understand that it is crucial to maintain consistent interpretation of our data throughout the manuscript. We will carefully reevaluate these sections and adjust our phrasing to ensure that our interpretations accurately reflect our results.

1. It is not clear why the authors investigated CD47 expression.

Thank you for your question regarding our investigation of CD47 expression. CD47, also known as integrin-associated protein, is ubiquitously expressed on many cell types, including red blood cells (RBCs). In the context of our study, we used CD47 expression as an indicator of young RBCs, as CD47 is known to be highly expressed on newly produced RBCs. Our intention was to use CD47 positive cells as a proxy for new RBC production, which would give us insights into erythropoiesis under hypobaric hypoxia conditions. This marker thus provides valuable information about the rate and effectiveness of erythropoietic response to hypoxic stress. However, according to others reviewers’ suggestion, we removed this part of results in the revised manuscript.

Minor:

1. Y axis is often labeled without sufficient detail.

1. The legends do not specify the exact statistical tests.

1. Some in vivo exp contain n=3 which is relatively low for mouse-based studies.

Some suggestions for the text:

Line 60: is the main cause of erythrocytosis which in turn alleviates..

62-66 - argumentation is not clear/grammatically correct and should be rephrased (eg, „RBC homeostasis is disturbed and never formed into a homeostasis status" - „homeostasis.. is never formed into a homeostasis status" sounds incorrect.

Ref # 8 - does not fit, I assume this was a mistake and the authors aimed to cite a Review article by Slusarczyk and Mleczko-Sanecka in Genes. However, this reference seems appropriate to be discussed in the Discussion section as it is very directly connected to the content of the present manuscript

76-78 - unclear/incomplete sentence (binding of iron to Tf and Tf-Fe delivery to the erythroid compartment is missing in this sentence, please, rephrase)

80 - iron is not stored ON FtL

90 - should be written: important role in iron recycling from RBCs

94 - phrasing 'damage of erythrophagocytosis' is incorrect

96-97 - should be written, for example: 'followed by eryptosis and iron recycling defects in the spleen'

282 - the sentence is grammatically incorrect and unclear.

292-94 - the statement is completely unclear, what can 'inhibit the excessive proliferation of RBCs'? What does it mean?

Reference to tuftsin was not provided (Am J Gastroenterol, 1999;94:391-397; PLoS One. 2012;7(4):e34933)

How quantification of microscopy images for F4/80 signal was performed?

In Figure 5, more explanation is required for the readers regarding the measured genes/proteins - why the patter of gene expression changes suggest ferroptosis?

Writing that ferroptosis INHIBITS phagocytosis is incorrect

Line 460 is unclear

468 - erythrocytophagy is not a commonly used term/

We are grateful for your keen eye and the time you have taken to provide such thorough feedback. It will undoubtedly help us to significantly enhance the clarity and completeness of our research. We have modified the corresponding sections in our manuscript to include these details. The comments have helped us ensure that our methodology is transparent and our findings are presented clearly. We have taken all your comments into consideration in our revision. we also have revised our manuscript to discuss these alternative interpretations more clearly and to acknowledge the potential limitations of our data.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: High-altitude polycythemia (HAPC) occurs in high-altitude (HA) environments and involves an imbalance between erythropoiesis and eryptosis. Spleen/splenic macrophages are an important primary tissue/cell of eryptosis and iron recycling. However, the role of the spleen in the pathogenesis of HAPC and the effect of hypobaric hypoxia (HH) on the biology of the spleen and splenic macrophages are still unclear. We used a mouse hypobaric hypoxia (HH) exposure model to simulate an in vivo study of 6000 m HA exposure. For in vitro studies, we used a primary splenic macrophage model treated with 1% hypoxia. We found that the HH-treated mouse model promoted erythropoiesis and led to erythrocytosis. In addition, HH exposure resulted in marked splenic contraction followed by splenomegaly for up to 14 days. HH exposure impaired the red blood cell (RBC) handling capacity of the spleen, which was caused by a decrease in splenic macrophages in the red pulp. Moreover, HH treatment for 7 and 14 days promoted iron mobilization and ferroptosis in the spleen, as reflected by the expression of metabolism-related proteins and ferroptosis-related proteins. All of the protein expression levels were similar to the gene expression levels in human peripheral blood mononuclear cells. Single-cell sequencing of the spleen further demonstrated a significant decrease in macrophages in the spleen 7 days after HH exposure. In in vitro studies, we confirmed that primary splenic macrophages decreased and induced ferroptosis following hypoxic treatment, which was reversed by pre-treatment with the ferroptosis inhibitor ferrostatin-1. Taken together, HH exposure induces splenic ferroptosis, especially in red pulp macrophages, which further inhibits the clearance of RBCs from the spleen. As such, it promotes the retention of RBCs in the spleen and causes splenomegaly, which may further lead to the persistent production of RBCs and ultimately to the development of HAPC.

Revised version: This study investigates the role of the spleen and splenic macrophages in the development of high-altitude polycythemia (HAPC), a condition characterized by an imbalance between erythropoiesis and eryptosis imbalance, typically observed in high-altitude (HA) environments. We employed a mouse model subjected to hypobaric hypoxia (HH) to simulate the conditions of a 6000 m HA exposure. For in vitro examination, primary splenic macrophages were treated with 1% hypoxia. Our findings revealed that the HH treatment enhanced erythropoiesis, resulting in erythrocytosis, with marked splenic contraction evident, later progressing to splenomegaly over a 14-day period. HH exposure also impaired the ability of the spleen to process red blood cells (RBCs), predominantly due to a decrease in splenic macrophages within the red pulp. Furthermore, the application of HH treatment over 7 and 14-day intervals resulted in increased iron mobilization and onset of ferroptosis within the spleen, as corroborated by the expression of iron metabolism-related and ferroptosis-related proteins. The expression levels of these proteins mirrored gene expression levels in human peripheral blood mononuclear cells. Subsequent single-cell sequencing of the spleen demonstrated a substantial decrease in macrophages 7 days post-HH exposure. In vitro investigations confirmed the decline in primary splenic macrophages and induction of ferroptosis following hypoxic treatment, which were reversed by pre-treatment with the ferroptosis inhibitor ferrostatin-1. In summary, the data suggested that HH exposure instigates splenic ferroptosis, predominantly in the red pulp, thereby hampering the RBCs clearance in the spleen. This leads to increased 46 RBCs retention within the spleen, triggering splenomegaly, which may potentially foster continuous RBCs production and accelerate HAPC progression. The major conclusion from this study elucidates the critical role of spleen and splenic macrophages in the pathogenesis of HAPC.

Background

Earlier version: A plateau is a special environment characterized by low atmospheric pressure and low partial oxygen pressure. Long-term exposure to plateau environments may lead to chronic mountain disease [1]. High-altitude polycythemia (HAPC) is a common and widespread chronic mountain sickness characterized by excessive erythrocytosis [2]. Hypobaric hypoxia (HH) is the main cause of erythrocytosis to alleviate hypoxic conditions in tissues under high-altitude (HA) exposure [3]. In a healthy organism, erythropoiesis in bone marrow is equivalent to eryptosis in the spleen [4]. Although HA/HH exposure can affect the balance between the formation and clearance of RBCs, a new physiological RBC steady state can be rapidly formed in normal individuals [5]. Nevertheless, RBC homeostasis is disturbed and never formed into a homeostasis status in HAPC patients, which finally results in a continuous increase in RBCs [4]. It is well known that HA/HH exposure may lead to erythropoiesis; however, the pathogenesis of eryptosis under HA/HH conditions remains unclear.

Revised version: A plateau is a special environment characterized by low atmospheric pressure and low partial oxygen pressure. Long-term exposure to plateau environments may lead to chronic mountain disease [1]. High-altitude polycythemia (HAPC) is a common and widespread chronic mountain sickness characterized by excessive erythrocytosis [2]. Hypobaric hypoxia (HH) is the main cause of erythrocytosis, which in turn alleviates hypoxic conditions in tissues under high-altitude (HA) exposure [3]. In a healthy organism, a balance is maintained between erythropoiesis in the bone marrow and eryptosis in the spleen [4]. Even though HA/HH exposure can disrupt the equilibrium between RBC formation and clearance, healthy individuals can swiftly establish a new physiological RBC steady state [5]. However, in patients with HAPC, RBCs homeostasis is disrupted, failing to reach a state of equilibrium, which ultimately leads to a persistent increase in RBCs [4]. It is well-established that exposure to HA/HH can induce erythropoiesis, yet the pathogenesis of eryptosis under these conditions remains poorly understood.

Earlier version: The spleen clears the altered RBCs in the circulation system, which helps to balance erythropoiesis and RBC clearance [6]. Moreover, red pulp macrophages (RPMs) of the spleen are the major scavengers that remove unhealthy, old, and malformed RBCs from the circulation to recover iron [7]. In brief, RPMs recognize the RBCs to be removed through the signal on their cell surface and trigger RBC endocytosis and digestion into heme in lysosomes [8]. Subsequently, the heme is decomposed by heme oxygenase-1 (HO-1) into biliverdin, carbon monoxide and iron [9]. Most of the iron recycling from heme is transported out of the cell through ferroportin (Fpn) and then participates in the formation of RBCs through the combination of transferrin (Tf) and transferrin receptor (TfR) on the cell membrane [10]. A portion of the released iron is loaded into cellular ferritin (Ft; including Ft-H and Ft-L). Ft-H oxidizes Fe2+ (ferrous ion) to Fe3+ (ferric ion) in the presence of O2, and then Fe3+ can be stored on Ft-L [11]. When the body’s iron metabolism is vigorous, nuclear receptor coactivator 4 (NCOA4) can recognize Ft-H, bring Ft into the lysosomal pathway for degradation, and release iron in Ft for iron recycling in the body [12]. However, some iron is free in cells in the form of ferrous iron (Fe2+), which is well known to be an important initiator of free radical oxidation. Moreover, the accumulation of large amounts of Fe2+ ions in cells may cause ferroptosis [13].

Revised version: The spleen plays a crucial role in maintaining erythropoietic homeostasis by effectively clearing impaired and senescent RBCs from circulation [6]. Particularly, red pulp macrophages (RPMs) within the spleen, serving as primary phagocytes, are responsible for clearing senescent, damaged, and abnormal erythrocytes from circulation to recycle iron [7]. RPMs initiate the process of RBC endocytosis and lysosomal digestion into heme, following the recognition of the RBCs earmarked for removal via their cell surface signals [8, 9]. Subsequently, heme oxygenase-1 (HO-1) decomposes the heme into biliverdin, carbon monoxide, and iron [10]. Most of the iron recycled from heme is transported out of the cell through a protein called ferroportin (Fpn). This iron then binds to transferrin (Tf), a plasma protein that transports iron in the blood. The iron-transferrin complex interacts with the transferrin receptor (TfR) on the cell membrane, contributing to the formation of RBCs in the erythroid compartment [11]. A portion of the released iron is loaded into cellular ferritin (Ft; including Ft-H and Ft-L). In the presence of oxygen, the Ft-H facilitates the oxidation of Fe2+ (ferrous ion) to Fe3+ (ferric ion). Subsequently, the Ft-L stores this Fe3+ [12]. The Ft-L features a nucleation site, consisting of a cluster of cavity-exposed carboxyl residues, which readily bind to Fe3+, thereby simplifying the storage process [13]. When the body’s iron metabolism is vigorous, nuclear receptor coactivator 4 (NCOA4) can recognize Ft-H, bring Ft into the lysosomal pathway for degradation, and release iron in Ft for iron utilization in the body [14]. However, some iron is free in cells in the form of ferrous iron (Fe2+), which is well known to be an important initiator of free radical oxidation. Moreover, the accumulation of large amounts of Fe2+ ions in cells may cause ferroptosis [15].

Earlier version: Although erythropoiesis under HA/HH is well studied thus far, the precise effects of HA/HH on erythrophagocytosis in the spleen remain largely unexplored. Based on the fact that the spleen is the major organ for RBC processing and that RPMs play an important role in recycling iron from RBC clearance, we investigated the effects of HA exposure on spleen/splenic macrophages in an HH-exposed mouse model (6000 m exposure conditions). In the current study, we further explored whether HA/HH could regulate the erythrocyte disposal and iron recycling of macrophages, and specifically predict the damage of erythrophagocytosis after HA/HH exposure. We demonstrated that HH exposure induced ferroptosis in the spleen, especially in macrophages, which led to a decrease in the number of macrophages, followed by eryptosis and iron recycling in the spleen. These findings may be clinically relevant to continuous pathological erythrocytosis and HAPC progression under HA exposure.

Revised version: Despite the extensive study of erythropoiesis under HA/HH conditions, the precise effects of HA/HH on erythrophagocytosis within the spleen remain largely unexplored. Considering the significant role of the spleen in RBC processing and the crucial function of RPMs in iron recycling from RBC clearance, we evaluated the impacts of HA exposure on the spleen/splenic macrophages using an HH-exposed mouse model (simulating 6000 m exposure conditions). In the current study, we sought to further investigate whether high-altitude/hypobaric hypoxia (HA/HH) can influence the erythrocyte disposal and iron recycling processes in macrophages. More specifically, we aimed to determine the potential impact of HA/HH exposure on the integrity of the erythrophagocytosis process. We discovered that exposure to HH triggered ferroptosis in the spleen, particularly in macrophages. This led to a reduction in macrophage numbers, which was subsequently followed by disruptions in eryptosis and iron recycling within the spleen. These findings may hold clinical significance, particularly in the context of continuous pathological erythrocytosis and the progression of HAPC under HA exposure.

Methods

HH mouse exposure model

Earlier version: Male C57BL/6 male mice (at 8 weeks of age) were obtained from the animal experimental centre of Nantong University. Mice were adapted to the facilities for 3 days before experiments. Mouse models of HH exposure were established by placing the mice in the decompression chamber (TOW-INT TECH, ProOx-810, China) for 1 d, 2 d, 3 d, 7 d, and 14 d. The parameters were set as follows: 6000 m above sea level, oxygen partial pressure of 11 kPa, lifting speed of 20 m/s, chamber pressure of 54 kPa, temperature of 25 ℃, humidity of 40%, and 12-h light/dark cycles. Tuftsin (1.5 mg/kg, MCE, HY-P0240, USA) was used to stimulate the phagocytosis of macrophages and was injected at 7 d and 11 d during the 14-d period of mice exposed to HH.

Revised version: Male C57BL/6 male mice (at 8 weeks of age) were obtained from the animal experimental centre of Nantong University. Mice were adapted to the facilities for 3 days before experiments. Mouse models of HH exposure were established by placing the mice in the decompression chamber (TOW-INT TECH, ProOx-810, China) for 1 d, 2 d, 3 d, 7 d, and 14 d. The parameters were set as follows: 6000 m above sea level, oxygen partial pressure of 11 kPa, lifting speed of 20 m/s, chamber pressure of 54 kPa, temperature of 25 ℃, humidity of 40%, and 12-h light/dark cycles. Tuftsin (1.5 mg/kg, MCE, HY-P0240, USA) was used to stimulate the phagocytosis of macrophages [16], and was injected at 7 d and 11 d during the 14-d period of mice exposed to HH.

Splenic macrophage culture and treatments

Earlier version: After isoflurane anaesthesia, mice were transcardially perfused with precooled PBS, and then splenic tissue (approximately 1mm3 in volume) was collected and separated in cold, phenol red-free Hanks’ balanced salt solution (HBSS; Gibco, Grand Island, NY, USA). The cell extracts were centrifuged at 500 rpm for 5 min at 4 ℃. The pellets were then resuspended and incubated with ammonium-chloride-potassium (ACK) lysing buffer (Thermo Fisher Scientific, USA) for 3 min at room temperature. Next, cell extracts were purified in a four-phase Percoll (Sigma‒Aldrich, St. Louis, MO, USA) gradient: 0, 30, 40, and 50% in HBSS. After centrifugation at 500 rpm for 20 min at 4°C, the macrophage-enriched fraction was collected at the interface between the 30% and 40% phases. Cells were washed in two rounds with HBSS and then resuspended in AIM V medium (Thermo Fisher Scientific). The cells in this medium were cultured at 37 ℃ and 5% CO2 for 12 h. After the cells were treated with ferrostatin-1 (Fer-1, 2 μm; Selleck Chemicals, S7243, USA) 1 h in advance, the cells were placed in a hypoxia workstation (Invivo2, Ruskinn, UK) for 24 h.

Revised version: After isoflurane anaesthesia, mice were transcardially perfused with precooled PBS, and then splenic tissue was collected. The spleen was then dissected into minuscule fragments, each approximately 1 mm3 in volume, and separated in cold, phenol red-free Hanks’ balanced salt solution (HBSS; Gibco, Grand Island, NY, USA). The cell extracts were centrifuged at 500 rpm for 5 min at 4 ℃. The pellets were then resuspended and incubated with ammonium-chloride-potassium (ACK) lysing buffer (Thermo Fisher Scientific, USA) for 3 min at room temperature. Next, cell extracts were purified in a four-phase Percoll (Sigma‒Aldrich, St. Louis, MO, USA) gradient: 0, 30, 40, and 50% in HBSS. After centrifugation at 500 rpm for 20 min at 4°C, the macrophage-enriched fraction was collected at the interface between the 30% and 40% phases. Cells were washed in two rounds with HBSS and then resuspended in AIM V medium (Thermo Fisher Scientific). The cells in this medium were cultured at 37 ℃ and 5% CO2 for 12 h. After the cells were treated with ferrostatin-1 (Fer-1, 2 μm; Selleck Chemicals, S7243, USA) 1 h in advance, the cells were placed in a hypoxia workstation (Invivo2, Ruskinn, UK) for 24 h.

Blood smear and hematological indices

The collected blood was gently mixed with EDTA (15 g/L) to obtain anticoagulant blood, which was detected by a hematology analyser (XE-5000, Sysmex Corporation, Japan). Blood smear preparation method: a blood sample is dropped onto a slide, and another slide is used to push the blood sample into a uniform thin layer at a constant speed. After the blood samples were dried, Wright’s stain solution was added and incubated for 1 min. Then, the slide was rinsed with water. Blood smears were photographed with a DM4000B microscope (Leica, Germany).

Added in the revised version: Lillie staining

Added in the revised version: The Lillie staining method was employed to ascertain Fe2+ deposition within the spleen, according to the protocol delineated by Liu et al [17]. Briefly, paraffin-embedded spleen sections underwent a process of dewaxing and rehydration via xylene and gradient alcohol. Then, the Lillie staining solution (G3320, Solarbio, China) was applied at 37°C for 50 min in a light-restricted environment, improved by a 5-min PBS wash. Subsequently, a mixture comprising 30 % hydrogen peroxide and methanol was incubated for 20 min, and then the sections were subjected to two 5-min washes with PBS. The DAB developing solution, which is not included in the Lillie bivalent iron staining kit, was applied for light-sensitive staining. The sections were stained with a nuclear fast red solution in a light-restricted setting for 10 minutes, followed by a 5-second wash with ddH2O. Lastly, the sections were dehydrated with gradient alcohol, cleared with xylene, and mounted with neutral resin.

Added in the revised version: Wright staining

Added in the revised version: Mice were anaesthetized and perfused with saline. The spleen was subsequently excised to prepare a single cell suspension, achieved by filtering through a 70 μm filter. Samples were centrifuged at 500 g for 5 min for cell separation. The cell pellet was subjected to three PBS washes. The suspension was then smeared onto a slide and air-dried. Following this, the slide was stained with Wright solution (G1040, Solarbio, China) for 3 min. An equivalent volume of pH 6.4 phosphate buffer was introduced, and the slide was gently agitated to allow mixing with the Wright’s stain solution for 5 min. After washing and drying, a microscope (DM4000B, Leica, Germany) was utilized for visualization and image acquisition.

Western blot

Tissue and cell samples were homogenized on ice after cell lysates and protease inhibitors were added. Centrifugation was performed at 12,000 rpm for 15 min at 4°C to collect the supernatant. Protein concentration was determined by the BCA detection method. A sample containing 30 μg protein was loaded and run in each well of SDS‒ PAGE gels. The membranes were incubated with the following antibodies: HO-1 (1:2000, Abcam, ab13243, USA); Ft-L (1:1000, Abcam, ab69090, USA); Ft-H (1:1000, Novus, NBP1-31944, USA); NCOA4 (1:1000, Santa Cruz, sc-373739, USA); TfR (1:1000, Thermo Fisher, 13-6800, USA); Fpn (1:1000, Novus, NBP1-21502, USA); ACSL4 (1:1000, Santa Cruz, sc-271800, USA); xCT (1:1000, Proteintech, 26864-1-AP, USA); Gpx4 (1:1000, Abcam, ab125066, USA); CD206 (1:1000, RD, AF2535, USA); CD16 (1:2000, RD, AF1960, USA); and β-actin (1:1000, Sigma‒Aldrich, A5316, USA). The secondary antibodies were as follows: goat anti-mouse (1:10000, Jackson, USA) and goat anti-rabbit (1:10000, Jackson, USA). The gray value of specific blots was scanned and analysed using ImageJ software (National Institute of Health, USA).

RT‒PCR analysis

Earlier version: Total RNA extraction and RT‒PCR analysis were performed essentially as described previously [14]. Primer sequences were as follows:

Revised version: Total RNA extraction and RT‒PCR analysis were performed essentially as described previously [18]. Primer sequences were as follows:

GEO analysis

Earlier version: GSE46480 samples were found in the GEO database of NCBI with the keyword “High Altitude” search. The data set was divided into two groups, one for the plain (Base Camp) and the other for the plateau (Altitude), with a sample size of 98 for each group. Blood samples were collected from the same person at McMurdo Station (48 m) and immediately transferred to Amundsen-Scott South Pole Station (2835 m) on the third day. R language was used for data analysis, and GraphPad Prism was used for statistics. Malondialdehyde (MDA), Cysteine (Cys) and Glutathione (GSH) content detection MDA was detected by a Micro Malondialdehyde Assay Kit (Solarbio, BC0025, China). Cys was detected by a Micro Cysteine Assay Kit (Solarbio, BC0185, China). GSH was detected by a Micro Reduced Glutathione Assay Kit (Solarbio, BC1175, China).

Revised version: GSE46480 samples were found in the GEO database of NCBI with the keyword “High Altitude” search. The data set was divided into two groups, one for the plain (Base Camp) and the other for the plateau (Altitude), with a sample size of 98 for each group. Blood samples were collected from the same person at McMurdo Station (48 m) and immediately transferred to Amundsen-Scott South Pole Station (2835 m) on the third day. R language was used for data analysis, and GraphPad Prism was used for statistics.

Added in the revised version: Malondialdehyde (MDA), Cysteine (Cys) and Glutathione (GSH) content detection

Added in the revised version: MDA was detected by a Micro Malondialdehyde Assay Kit (Solarbio, BC0025, China). Cys was detected by a Micro Cysteine Assay Kit (Solarbio, BC0185, China). GSH was detected by a Micro Reduced Glutathione Assay Kit (Solarbio, BC1175, China).

Immunofluorescence

Earlier version: Frozen sections of the spleen (9 μm) were stored at −80°C. Sections were placed at room temperature for 30 min before immunofluorescence staining, incubated in 10% BSA-PBS for 10 min, incubated overnight with F4/80-PE (1:200, Biolegend, 123110, USA), washed with 0.05% PBST three times, incubated with DAPI for 10 min, washed with PBS three times, sealed with 50% glycerine-PBS, and photographed with a confocal laser scanning microscope (SP8, Leica Microsystems, Wetzlar, Germany).

Revised version: Spleen sections of 9 μm thickness were flash-frozen and stored at –80°C. Before immunofluorescence staining, these sections were allowed to acclimate to room temperature for 30 minutes, then incubated in 10% BSA-PBS for 10 minutes. Next, the sections were incubated overnight with F4/80-PE (1:200, Biolegend, 123110, USA) and then washed thrice with PBS. Due to the limitations of F4/80 as a definitive macrophage marker, we concurrently conducted immunohistochemical analysis for heme oxygenase-1 (HO-1) (1:200, Abcam, ab13243, USA). Following this, sections were washed thrice with 0.05% PBST, incubated with secondary antibodies DAR-Alexa555 at room temperature for 2 hours, washed thrice with PBS again, and finally sealed with 50% glycerine-PBS. Fluorescence microscopy images were captured using a confocal laser scanning microscope (SP8, Leica Microsystems, Wetzlar, Germany). To quantify the immunostaining intensity of F4/80 and HO-1, we utilized ImageJ software. Specifically, we captured multiple fields of view for each slide, and the software was used to measure the mean intensity of the fluorescent signal in these images. These measurements were then normalized to the NN group to account for any experimental variations.

Tissue iron staining (DAB-enhanced Perls’ staining)

Earlier version: DAB-enhanced Perls’ staining for the spleen paraffin sections was performed as described previously [15]. Briefly, the sections were washed with PBS and cultured in freshly prepared Perls’ solution (1% potassium ferricyanide in 0.1 M hydrochloric acid buffer). The slides were then immersed and stained with DAB. All slides were counterstained with hematoxylin and visualized under a DM4000B microscope (Leica, Germany). Data were collected from three fields of view per mouse and semiquantitatively analysed with ImageJ software. Quantitative results of iron staining were finally normalized to the NN control group.

Revised version: DAB-enhanced Perls’ staining for the spleen paraffin sections was performed as described previously [19]. Briefly, the sections were washed with PBS and cultured in freshly prepared Perls’ solution (1% potassium ferricyanide in 0.1 M hydrochloric acid buffer). The slides were then immersed and stained with DAB. All slides were counterstained with hematoxylin and visualized under a DM4000B microscope (Leica, Germany). Data were collected from three fields of view per mouse and semiquantitatively analysed with ImageJ software. Quantitative results of iron staining were finally normalized to the NN control group.

Flow cytometry

Earlier version: 1) The steps of reticulocyte ratio detection were as follows: The experimental tube and negative control tube were prepared, 125 μL normal saline, 4 μL anticoagulant and 125 μL TO working solution (1 μg/mL, Sigma, 390062, USA) were added to each tube, and 250 μL normal saline and 4 μL anticoagulant were added to each tube of the negative control tube. After incubation at room temperature for 1 h, flow cytometry analysis was carried out by using the FL1 (488 nm/525 nm) channel.

Revised version: The steps of reticulocyte ratio detection were as follows: Whole blood was extracted from the mice and collected into an anticoagulant tube, which was then set aside for subsequent thiazole orange (TO) staining [20]. The experimental tube and negative control tube were prepared, 125 μL normal saline, 4 μL anticoagulant and 125 μL TO working solution (1 μg/mL, Sigma, 390062, USA) were added to each tube, and 250 μL normal saline and 4 μL anticoagulant were added to each tube of the negative control tube. After incubation at room temperature for 1 h, flow cytometry analysis was carried out by using the FL1 (488 nm/525 nm) channel.

Earlier version: 2) The steps for the determination of intracellular divalent iron content and lipid peroxidation level were as follows: The spleen was made into a single cell suspension, and then the RBCs were lysed. This step can be omitted from cell samples. A total of 1 × 106 cells were incubated with 100 μL of BioTracker Far-red Labile Fe2+ Dye (1 mM, Sigma, SCT037, USA) for 1 h or C11-Bodipy 581/591 (10 μM, Thermo Fisher, D3861, USA) for 30 min. After the cells were washed with PBS twice, flow cytometry analysis was carried out by using the FL6 (638 nm/660 nm) channel for determination of intracellular divalent iron content or the FL1 (488 nm/525 nm) channel for determination of lipid peroxidation level.

Revised version: The steps for the determination of intracellular divalent iron content and lipid peroxidation level were as follows: Splenic tissue was procured from the mice and subsequently processed into a single-cell suspension using a 40 μm filter. The RBCs within the entire sample were subsequently lysed and eliminated, and the remaining cell suspension was resuspended in PBS in preparation for ensuing analyses. A total of 1 × 106 cells were incubated with 100 μL of BioTracker Far-red Labile Fe2+ Dye (1 mM, Sigma, SCT037, USA) for 1 h or C11-Bodipy 581/591 (10 μM, Thermo Fisher, D3861, USA) for 30 min. After the cells were washed with PBS twice, flow cytometry analysis was carried out by using the FL6 (638 nm/660 nm) channel for determination of intracellular divalent iron content or the FL1 (488 nm/525 nm) channel for determination of lipid peroxidation level.

Earlier version: 3) The steps for detecting the mortality of spleen cells and peritoneal macrophages were as follows: 1 × 106 cells were incubated at room temperature for 40 min with 2 μM Calcein AM and 8 μM PI. Flow cytometry analysis was carried out by using FL1 (488 nm/525 nm, Calcein AM) and FL3 (488 nm/620 nm, PI) channels.

Revised version: The steps for detecting the mortality of spleen cells and peritoneal macrophages were as follows: 1 × 106 cells were incubated at room temperature for 40 min with 2 μM Calcein AM and 8 μM PI. Flow cytometry analysis was carried out by using FL1 (488 nm/525 nm, Calcein AM) and FL3 (488 nm/620 nm, PI) channels.

Earlier version: 4) The steps for detecting the number of M1/M2 macrophages in the spleen were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE (1:200), CD86-PE/Cyanine7 (1:20, Biolegend, 105014, USA), and CD206-Alexa 647 (1:80, BD, 565250) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80-PE), FL5 (488 nm/755 nm, CD86-PE/Cyanine7) and FL6 (638 nm/660 nm, CD206-Alexa 647) channels.

Revised version: The steps for detecting the number of M1/M2 macrophages in the spleen were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE (1:200), CD86-PE/Cyanine7 (1:20, Biolegend, 105014, USA), and CD206-Alexa 647 (1:80, BD, 565250) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80-PE), FL5 (488 nm/755 nm, CD86-PE/Cyanine7) and FL6 (638 nm/660 nm, CD206-Alexa 647) channels.

Earlier version: 5) The steps for detecting the number of monocytes in blood, spleen and bone marrow were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE, CD11b-PE/CY7 (1:2000, BD, 552850, USA), and Ly6C-APC (1:2000, Thermo Fisher, 17-5932-82) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80-PE), FL5 (488 nm/755 nm, CD11b-PE/CY7) and FL6 (638 nm/660 nm, Ly6C-APC) channels.

Revised version: The steps for detecting the number of monocytes in blood, spleen and bone marrow were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE, CD11b-PE/CY7 (1:2000, BD, 552850, USA), and Ly6C-APC (1:2000, Thermo Fisher, 17-5932-82) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80-PE), FL5 (488 nm/755 nm, CD11b-PE/CY7) and FL6 (638 nm/660 nm, Ly6C-APC) channels.

Phagocytosis of E. coli and RBCs

Earlier version: 1. E. coli was labelled with Cy5.5 (5 mg/mL), and RBCs were labelled with NHS-biotin (20 mg/mL). Macrophages (1 × 106) were coincubated with E. coli-Cy5.5 or RBC- Biotin for 30 min and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL6 (638 nm/660 nm) to determine the phagocytosis of spleen by E. coli. Macrophages (coincubated with Biotin-RBCs) were incubated with streptavidin-FITC for 3 h and washed twice with 2% mouse serum PBS. FL1 (488 nm/525 nm) was used for flow cytometry analysis to determine the phagocytosis of RBCs in the spleen.

Revised version: E. coli was labelled with Cy5.5 (5 mg/mL), and RBCs were labelled with NHS-biotin (20 mg/mL). Macrophages (1 × 106) were coincubated with E. coli-Cy5.5 or RBC-Biotin for 30 min and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL6 (638 nm/660 nm) to determine the phagocytosis of spleen by E. coli. Macrophages (coincubated with Biotin-RBCs) were incubated with streptavidin-FITC for 3 h and washed twice with 2% mouse serum PBS. FL1 (488 nm/525 nm) was used for flow cytometry analysis to determine the phagocytosis of RBCs in the spleen.

Single-cell RNA Sequencing

The spleen tissues were surgically removed and stored in MACS Tissue Storage Solution (Miltenyi Biotec, Bergisch Gladbach, Germany) until processing. Single-cell dissociation was performed by the experimentalists at the GENECHEM laboratory (Shanghai, China). Dissociated single cells were then stained for viability assessment using Calcein-AM (BD Biosciences, USA) and Draq7 (BD Biosciences). The BD Rhapsody system was used to capture transcriptomic information from single cells. Single-cell capture was achieved by random distribution of a single-cell suspension across >200,000 microwells through a limited dilution approach. Beads with oligonucleotide barcodes were added to saturation so that a bead was paired with a cell in a microwell. The cells were lysed in the microwell to hybridize mRNA molecules to barcoded capture oligos on the beads. Beads were collected into a single tube for reverse transcription and ExoI digestion. Upon cDNA synthesis, each cDNA molecule was tagged on the 5′ end (that is, the 3′ end of an mRNA transcript) with a unique molecular identifier (UMI) and cell barcode indicating its cell of origin. Whole transcriptome libraries were prepared using the BD Rhapsody single-cell whole-transcriptome amplification (WTA) workflow, including random priming and extension (RPE), RPE amplification PCR and WTA index PCR. The libraries were quantified using a High Sensitivity D1000 ScreenTape (Agilent) and High Sensitivity D1000 Reagents (Agilent) on a 4150 TapeStation System (Agilent, Palo Alto, CA, USA) and the Qubit High Sensitivity DNA assay (Thermo Fisher Scientific). Sequencing was performed on an Illumina sequencer (Illumina Nova Seq 6000, San Diego, CA) in a 150 bp paired-end run.

Statistical analysis

Earlier version: Statistical analysis was performed with GraphPad Prism 8.0 software. Two-tailed Student’s t test or one-way analysis of variance (ANOVA) with Tukey’s post hoc test. P < 0.05 was considered statistically significant.

Revised version: Statistical analysis was conducted using GraphPad Prism 8.0 software. All values were represented as the mean ± standard error (SEM). The homogeneity of variance was verified before the application of a parametric test. A two-tailed Student’s t-test was employed for data from two groups, while a one-way analysis of variance (ANOVA) with multiple comparisons using Tukey’s post hoc test was utilized for data from multiple groups. A p-value of less than 0.05 was considered statistically significant.

Results

HH exposure promotes erythrocytosis in mice

Earlier version: To mimic 6000 m HA exposure, we placed C58BL/6 mice in an animal hypobaric oxygen chamber and detected the blood indices in the blood of the mice after HH exposure for different times. The blood smear showed that the number of RBCs was increased from 3 to 14 days after HH exposure compared with the normobaric normoxia (NN) group (Figure 1A). Routine blood tests further confirmed the results in blood smears, which showed that the RBC number (Figure 1B), HGB content (Figure 1C) and HCV value (Figure 1D) were all increased significantly to varying degrees, while MCH was not changed after HH exposure (Figure 1E). We further performed flow cytometry by TO staining to detect reticulocytes after 7 and 14 days of HH treatment. The results showed that at both HH 7 and 14 d of exposure, the reticulocyte proportion was increased (Figure 1F). In addition, the number of CD47+ RBCs in the blood was also significantly increased at both 7 (Figure 1G and H) and 14 days (Figure 1G and I) after HH exposure. These results suggested that the HH-treated mouse model mimics HA exposure well, which promotes erythropoiesis and results in erythrocytosis in mice following HH exposure.

Revised version: To mimic 6000 m HA exposure, we placed C58BL/6 mice in an animal hypobaric oxygen chamber and detected the blood indices in the blood of the mice after HH exposure for different times. The blood smear showed that the number of RBCs was increased from 3 to 14 days after HH exposure compared with the normobaric normoxia (NN) group (Figure 1A). Routine blood tests further confirmed the results in blood smears, which showed that the RBC number (Figure 1B), HGB content (Figure 1C) and HCT value (Figure 1D) were all increased significantly to varying degrees, while MCH was not changed after HH exposure (Figure 1E). We further performed flow cytometry by TO staining to detect reticulocytes after 7 and 14 days of HH treatment. The results showed that at both HH 7 and 14 d of exposure, the reticulocyte proportion was increased (Figure 1F). These results suggested that the HH-treated mouse model mimics HA exposure well, which promotes erythropoiesis and results in erythrocytosis in mice following HH exposure.

Added in the revised version: HH exposure promotes erythrocytosis in mice.

Earlier version: HH exposure promotes erythrocytosis in mice. The C57BL/6 mice were treated with normobaric normoxia (NN) and hypobaric hypoxia (HH) for 3, 7 and 14 days, and blood was collected for the following tests. (A) Observation of the number and morphology of RBCs by blood cell smear. Routine blood analysis of RBC counts (B) and HGB (C), HCT (D), and MCH content (E). (F) Representative dot plots (one of three experiments) of TO-positive cells (reticulocytes) in total populations of spleen. (G) CD47 staining in splenic cells was analysed by flow cytometry. Representative dot plots (one of three experiments) of CD47 expression after HH for 7 (H) and 14 (I) days in total populations of the spleen. Data were expressed as the means ± SEM (n = 5 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: The C57BL/6 mice were treated with normobaric normoxia (NN) and hypobaric hypoxia (HH) for 3, 7 and 14 days, and blood was collected for the following tests. (A) Analysis of the number and morphology of RBCs via a blood cell smear. Routine blood analysis was performed for RBC counts (B), as well as HGB (C), HCT (D), and MCH content (E). (F) Representative dot plots (from three experiments) indicate TO-positive cells (reticulocytes) in whole blood. Data are expressed as means ± SEM (n = 5 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Spleen inhibits the immoderate increase in RBCs under HH conditions

Earlier version: To determine the roles of the spleen in RBC homeostasis under HA/HH, we investigated the effects of HH on the morphology, volume and weight of the spleen as well as erythrocyte indices. As shown in Figure 2, the spleen volume and weight were decreased significantly after HH exposure for 1 day compared to NN treatment (Figure 2A-C). However, the spleen was obviously enlarged from 2 to 14 days after HH exposure (Figure 2A-C). The results indicated that the spleen contracted, the stored RBCs in the spleen were released into the blood at 1 day, and the RBCs were produced and/or retained in the spleen from 2 to 14 days after HH exposure. In addition, we also investigated the effect of the spleen on RBC homeostasis under HH conditions and observed whether its clearance effect on RBCs in HH was compensated by the liver or other mononuclear macrophage systems. We removed the spleen from mice using splenectomies and then performed HH exposure for 14 days to detect the RBC counts and blood deposition. The results showed that compared with the splenectomy group mice under NN conditions or the sham group mice exposed to HH, erythrocyte deposition (Figure 2D) and counts (Figure 2E), HGB (Figure 2F) and HCV (Figure 2G) levels were all significantly increased 14 days after HH splenectomy, while MCH did not change (Figure 2H). However, these indices were not changed in the mice with or without splenectomies under NN conditions (Figure 2D and E). It is suggested that the liver and other mononuclear macrophages did not play an effective compensatory role under HH exposure conditions. In addition, the spleen is essential to maintain the homeostasis of RBCs after HH exposure, which can inhibit the excessive proliferation of RBCs after HH exposure.

Revised version: To determine the roles of the spleen in RBC homeostasis under HA/HH, we investigated the effects of HH on the morphology, volume and weight of the spleen as well as erythrocyte indices. As shown in Figure 2, the spleen volume and weight were decreased significantly after HH exposure for 1 day compared to NN treatment (Figure 2, A-C). However, the spleen was obviously enlarged from 2 to 14 days after HH exposure (Figure 2, A-C). The results indicated that the spleen contracted, the stored RBCs in the spleen were released into the blood at 1 day, and the RBCs were produced and/or retained in the spleen from 2 to 14 days after HH exposure. In our research, we also examined the influence of the spleen on RBCs homeostasis under HH conditions. We investigated whether the role of spleen in RBCs clearance under HH conditions could be compensated by the liver or other components of the mononuclear macrophage system. To conduct this, we performed splenectomies on mice and subsequently exposed them to HH conditions for 14 days. This allowed us to monitor RBCs counts and blood deposition. Our findings indicated that, in comparison to both the splenectomized mice under NN conditions and the sham-operated mice exposed to HH, erythrocyte deposition (Figure 2D) and counts (Figure 2E), as well as HGB (Figure 2F) and HCT (Figure 2G) levels, significantly increased 14 days post-splenectomy under HH conditions. Meanwhile, MCH levels remained stable (Figure 2H). These indices did not vary in mice, regardless of whether they had undergone a splenectomy, under NN conditions (Figure 2D and E). These results suggest that the liver and other mononuclear macrophages do not effectively compensate under HH conditions. Furthermore, our research highlights the crucial role of the spleen in maintaining RBCs homeostasis under conditions of HH. The study suggests that the spleen functions to regulate the proliferation of RBCs, preventing their overproduction in response to HH exposure.

Added in the revised version: The spleen plays an important role in suppressing the immoderate increase in RBCs under HH conditions.

Earlier version: The spleen plays an important role in suppressing the immoderate increase in RBCs under HH conditions. C57BL/6 mice with or without splenectomy were treated with NN and HH for different numbers of days, and the spleen and blood were collected for subsequent analyses. (A) Morphological observation, (B) spleen volume and (C) spleen weight detection after HH exposure. (D-H) Blood observation and hematological index detection. Data were expressed as the means ± SEM (n = 9 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: C57BL/6 mice with or without splenectomy were treated with NN and HH for varying durations, and the spleen and blood were collected for subsequent analyses. (A) Morphological observation, (B) spleen volume, and (C) spleen weight were determined post-HH exposure. (D-H) Blood observation and hematological index detection followed. Data are expressed as the means ± SEM (n = 9 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Added in the revised version: Reduction of splenic macrophages due to HH exposure

Removed in the revised version: HH exposure causes a decrease in the number of macrophages in spleen

Earlier version: Since macrophages are the main cell population for processing RBCs in the spleen under physiological conditions, we next investigated the counts and activity of macrophages after 7 or 14 days of HH exposure via single-cell sequencing and flow cytometry. Single-cell sequencing of the spleen showed that splenic macrophages were significantly decreased after 7 days of HH exposure (Figure 3A-C). Moreover, the flow cytometry of Calcein/PI double staining showed that the living cells of spleen were significantly decreased while the dead cells were significantly increased after HH exposure for 7 and 14 days (Figure 3D-F). We further found that both CD16 (M1 macrophage marker) and CD206 (M2 macrophage marker) expression were decreased at both 7 and 14 days of HH exposure (Figure 3G-L). The flow cytometry results were also consistent with the Western blot results, which showed that mature macrophages (F4/80+/CD11b+), including M1-type (F4/80+/CD86+) and M2-type (F4/80+/CD206+) macrophages, in the spleens of mice decreased after HH exposure for 14 days (Figure 3M). Next, to identify the migration and differentiation of monocytes from bone marrow to spleen, we further analysed the expression of chemokines CCL2, CCL7, Csf1 and Csf2 in spleen by qPCR and detected the number of monocytes in bone marrow and spleen by flow cytometry. The results showed that CCL2 and CCL7 (Figure 3N) and Csf1 and Csf2 (Figure 3O) expression were significantly decreased in the spleen after HH exposure at 7 and 14 days. Moreover, monocytes (Ly6C+/CD11b+) in the bone marrow (Figure 3P-R) and spleen (Figure 3P and S-T) also declined after HH for 7 and 14 days. Then, we assessed macrophage depletion in the spleen under HH via detection of the expression and distribution of HO-1 and F4/80. Figure 3U shows that both HO-1 and F4/80 were distributed in the red pulp of the spleen, and their expression was also decreased after HH exposure. These results suggested that the number of splenic macrophages decreased after HH exposure, resulting in the impairment of the processing ability of splenic erythrocytes.

Revised version: Considering macrophages as the primary cell type responsible for processing RBCs within the spleen under physiological conditions, we subsequently investigated the population and activity of these macrophages after exposure to HH for 7 or 14 days, employing flow cytometry and single-cell sequencing techniques. Calcein/PI double staining, examined via flow cytometry, revealed a significant decrease in viable spleen cells and a concurrent increase in dead cells following 7 and 14 days of HH exposure (Figure 3, A-C). Single-cell sequencing further disclosed a significant reduction in splenic macrophages after 7 days of HH exposure (Figure 3, D-F). Additionally, the expressions of CD16 (M1 macrophage marker) and CD206 (M2 macrophage marker) were observed to decrease after both 7 and 14 days of HH exposure (Figure 3, G-L). Western blot analysis corroborated the flow cytometry findings, showing a decrease in mature macrophages (F4/80+/CD11b+), including both M1-type (F4/80+/CD86+) and M2-type (F4/80+/CD206+) macrophages, in the spleens of mice after 7 days of HH exposure (Figure 3M). To further elucidate the migration and differentiation of monocytes from the bone marrow to the spleen, we analyzed the expression of chemokines CCL2, CCL7, Csf1, and Csf2 in the spleen using qPCR and determined the number of monocytes in the bone marrow (BM) and spleen via flow cytometry. Our results indicated a significant reduction in the expression of CCL2, CCL7 (Figure 3N), Csf1, and Csf2 (Figure 3O) in the spleen after 7 and 14 days of HH exposure. Furthermore, the number of monocytes (Ly6C+/CD11b+) in the bone marrow (Figure 3, P-R) and spleen (Figure 3, P and S-T) also exhibited a decline after 7 and 14 days of HH exposure. 373 We evaluated the depletion of macrophages in the spleen under HH conditions by examining the expression and distribution of HO-1 and F4/80. Figures 3U-W depict a decrease in both HO-1 and F4/80, predominantly within the spleen’s red pulp following HH exposure. Together, these findings indicate a reduction in the number of splenic macrophages after HH exposure, which could impair the spleen’s capacity to process erythrocytes.

Added in the revised version: HH exposure results in a decrease in the number of splenic macrophages.

Earlier version: HH exposure causes a decrease in the number of macrophages in spleen. The C57BL/6 mice were treated with NN and HH for 7 or 14 days, and the spleen and blood were collected for subsequent detection. (A) Visualization of spleen cell clusters using uniform manifold approximation and projection (UMAP). Each color represents a cluster of specific cells of the spleen, which was identified by a defined gene expression profile. (B) The proportions of different macrophage clusters derived from either NN-or HH-treated mouse spleens were analysed. (C) Individual macrophages derived from either NN-or HH-treated mouse spleens are denoted. (D) Calcein/PI double staining was analysed by flow cytometry. (E and F) Cell death proportions in total splenic cells are given as bar graphs. (G and J) Western blot detecting CD16 and CD206 protein expression in the spleen. (H-I and K-L) Statistical analysis of CD16 and CD206 expression in H and K. (M) F4/80 staining with CD11b, CD86 or CD206 in splenic cells was analysed by flow cytometry. (N) qPCR analysis of CCL2 and CCL7 expression in the spleen. (O) qPCR analysis of Csf1 and Csf2 expression in the spleen. (P) Flow cytometry detection of CD111b and Ly6C double-stained cells in the BM and spleen. CD11bhiLy6Chi cell proportions in the BM (Q-R) and spleen (S-T) are given as bar graphs. (U) HO-1 and F4/80 expression in the spleen after HH exposure for 7 and 14 days. Data were expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P< 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: C57BL/6 mice were treated with NN and HH for 7 or 14 days, and the spleen and blood were collected for subsequent detection. (A) Calcein/PI double staining was analysed by flow cytometry. (B and C) The bar graphs represent cell death proportions in total splenic cells. (D) Uniform manifold approximation and projection (UMAP) visualized spleen cell clusters. Each color represents a cluster of specific spleen cells, identified by a distinct gene expression profile. (E) The proportions of different macrophage clusters derived from either NN-or HH-treated mouse spleens were analysed. (F) Individual macrophages derived from either NN-or HH-treated mouse spleens are denoted. (G and J) Western blot detecting CD16 and CD206 protein expression in the spleen. (H-I and K-L) Statistical analysis of CD16 and CD206 expression in H and K. (M) F4/80 staining with CD11b, CD86 or CD206 in splenic cells was analysed by flow cytometry. (N) qPCR analysis of CCL2 and CCL7 expression in the spleen. (O) qPCR analysis of Csf1 and Csf2 expression in the spleen. (P) Flow cytometry detection of CD11b and Ly6C double-stained cells in the BM and spleen. CD11bhiLy6Chi cell proportions in the BM (Q-R) and spleen (S-T) are given as bar graphs. (U) HO-1 and F4/80 expression in the spleen after 0 (NN), 7, and 14 914 days of HH exposure. (V) The relative fluorescence of HO-1 and (W) F4/80 915 expression in the spleen as described in (U) was quantified. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Earlier version: HH exposure decreases erythrophagocytosis and the iron processing capacity of the spleen

Revised version: The impact of HH exposure on erythrophagocytosis and iron processing capacity of the spleen

Earlier version: We further investigated the effects of HH exposure on erythrocyte phagocytosis and heme iron recycling in splenic macrophages. We used Cy5.5 dye to label E. coli (Figure 4A) and incubated them with splenic cells from HH-exposed mice in vitro, and the phagocytosis of splenic cells was detected using flow cytometry (Figure 4B). Figure 4B shows that the phagocytic ability of splenic cells in vitro was significantly decreased after 7 days of HH exposure. In addition, we used biotin to label autologous RBCs and injected them into mice followed by HH treatment in vivo (Figure 4C and D), and the phagocytosis of splenic cells was detected using flow cytometry. At 7 and 14 days after HH, the number of biotin-labelled RBCs in blood and spleen cells decreased significantly (Figure 4E-H). Although the labelled RBC decay declined automatically with time, the degree of RBC decay in the HH exposure group was weaker than that in the NN group in both the blood (Figure 4E and G) and spleen (Figure 4F and H). Compared with blood, the phagocytic ability of mouse spleen macrophages to RBCs decreased more significantly after HH exposure, especially on the 14th day (Figure 4G and H). Furthermore, we injected Tuftsin, which can stimulate phagocytosis of macrophages, and performed immunofluorescence and iron staining to study the heme iron recycling ability of the spleen under HH conditions (Figure 4I). As shown in Figure 4J, compared with the NN group, F4/80 expression was decreased significantly as well as Fe3+ deposition in the red pulp after HH exposure. However, Tuftsin induced F4/80 expression and Fe3+ deposition in the red pulp of the spleen under HH (Figure 4K and L). These results suggested that the erythrocyte phagocytosis and heme iron recycling ability of the spleen were suppressed, which was mainly due to the decrease in macrophages in red pulp under HH conditions.

Revised version: We investigated the influence of HH exposure on erythrocyte phagocytosis and heme iron recycling within splenic macrophages. We administered NHS-biotin intravenously to the mice followed by HH treatment in vivo. The clearance of RBCs was monitored by detecting the retention of biotin-labelled RBCs using flow cytometry in both the blood and spleen of the HH-treated mice. We observed a substantial decrease in biotin-labelled RBCs in the blood and spleen cells at 7-and 14-days post HH exposure (Figure 4, A-D). Despite the natural decrease of labelled RBC over time, the decay was more pronounced in the HH exposure group compared to the normobaric normoxia (NN) group in both the blood (Figure 4, A and C) and spleen (Figure 4, B and D). The phagocytic ability of mouse spleen macrophages towards RBCs was significantly more impaired after HH exposure, particularly on the 14th day (Figure 4, A-D). Additionally, to substantiate our findings, we quantified the retention of RBC counts by Wright staining in splenic cells after 7 and 14 days of HH exposure (Figure 4, E-F). This confirmed impaired erythrophagocytosis in the spleen under HH exposure, as evidenced by an increase in deformation and retention of RBCs following 7 and 14 days of HH exposure. Furthermore, we employed Tuftsin to stimulate phagocytosis of macrophages and carried out immunofluorescence and iron staining to examine the heme iron recycling capability of the spleen under HH conditions (Figure 4G). As depicted in Figure 4H, compared with the NN group, F4/80 expression and iron deposition in the red pulp were significantly reduced after HH exposure. However, Tuftsin led to an induction of F4/80 expression and iron deposition in the red pulp of the spleen under HH (Figure 4, I-L). These results imply that the capacity for erythrocyte phagocytosis and heme iron recycling by the spleen was compromised, principally due to the reduction in macrophages in the red pulp under HH conditions.

Added in the revised version: HH exposure decreases erythrophagocytosis and iron processing capacity in the spleen.

Earlier version: HH exposure decreases erythrophagocytosis and iron processing capacity in the spleen. (A) Ecoil was labelled with Cy5.5 dye and photographed under a fluorescence microscope. (B) The phagocytic activity of splenic cells for Cy5.5-labelled Ecoil in vitro. (C) Autologous erythrocytes were labelled with biotin followed by NIH-FITC detection and photographed under a fluorescence microscope. (D) Experimental diagram of phagocytic activity in splenic cells after HH exposure in vivo. Flow cytometry detection of FITC-stained cells in the blood (E) and spleen (F) after HH exposure for 7 and 14 days. The FITC-positive cell proportions in the blood (G) and spleen (H) are given as bar graphs. (I) Experimental design of C57Bl6 mice (n=3 per group) treated at days 7 and 11 with a single intravenous injection of Tuftsin (1.5 mg/kg), followed by HH exposure until day 14 for isolation of the spleen for F4/80 immunohistochemistry and iron staining. (J) F4/80 and iron staining in the spleen after HH exposure for 14 days with Tuftsin treatment. (K) Quantitative results of F4/80 expression in the spleen described in (J). (L) Semiquantitative iron levels in the spleen described in (J). Data were expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: Flow cytometry revealed the proportion of FITC-stained cells in the blood (A) and spleen (B) after 7 and 14 days of HH exposure. The proportions of these FITC-positive cells in the blood (C) and spleen (D) are presented as bar graphs. (E) Wright staining of splenic cells after HH exposure for 0 (NN), 7 and 14 days. (F) The proportion of retention RBCs within the spleen cells. (G) Illustrating the experimental design, in which C57BL6 mice (n=3 per group) were treated on days 7 and 11 with a single intravenous Tuftsin injection (1.5 mg/kg), followed by HH exposure until day 14 for spleen isolation and subsequent F4/80 immunohistochemistry and iron staining. (H) F4/80 and iron staining in the spleen after HH exposure for 14 days with Tuftsin treatment. (I) Quantitative the relative fluorescence of F4/80 expression in the spleen described in (H). (J) Semiquantitative iron levels in the spleen described in (H). Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Added in the revised version: HH exposure induces iron mobilization and ferroptosis in the spleen

Added in the revised version: To elucidate the precise mechanisms underlying the macrophage reduction caused by HH, we examined the expression of proteins related to iron metabolism and ferroptosis in mice treated with HH and analyzed the corresponding gene expression in peripheral blood mononuclear cells (PBMCs) from healthy humans acutely exposed to HA conditions using GEO data (No. GSE46480). The results showed that, compared to the NN group, the expression levels of HO-1, Ft-L, Ft-H, NCOA4, and xCT were decreased, while Fpn, TfR, and ACSL4 expressions were significantly increased after 7 and 14 days of HH exposure (Figure 5, A-B, D-E; Figure S1, A-D).

HH exposure promotes iron mobilization and induces ferroptosis in the spleen.

Removed in the revised version: To clarify the precise mechanisms of macrophage reduction caused by HH, we examined the expression of iron metabolism-and ferroptosis-related proteins in mice treated with HH and analysed the expression of the corresponding genes in peripheral blood mononuclear cells (PBMCs) of healthy people acutely exposed to an HA environment from GEO data (No. GSE46480). The results showed that compared to the NN group, the expression levels of HO-1, Ft-L, Ft-H, NCOA4, and xCT were decreased, while the expression levels of Fpn, TfR and ACSL4 were significantly increased both at 7 and 14 days following HH exposure (Figure 5A-B, D-E; Figure S1, A-D). Except for NCOA4, the expression changes in iron metabolism and ferroptosis genes in PBMCs were consistent with our WB results (Figure 5C and F). There were no significant changes in the expression of the GPX4 gene and protein in human PBMCs and mouse spleens. The changes in iron metabolism-and ferroptosis-related genes in PBMCs can fully reflect the protein changes in the spleen under HH exposure. The levels of Fe2+ and lipid ROS in the spleen were detected by flow cytometry, and the contents of MDA, GSH and Cys were detected by biochemical detection kits. As shown in Figure 5G and H; Figure S1I-K, the levels of Fe2+ (Figure 5G; Figure S1E and F) and lipid ROS (Figure 5H; Figure S1G and H) in the spleen were significantly increased after HH exposure. In addition, the content of MDA (Figure 5I; Figure S1I) in the spleen was increased, while Cys (Figure 5J; Figure S1J) and GSH (Figure 5K; Figure S1K) were decreased significantly after HH exposure. All these results suggested that iron mobilization in the spleen was enhanced and ferroptosis was then induced after 7 days of HH exposure.

Earlier version: HH exposure promotes iron mobilization and induces ferroptosis in the spleen. The C57BL/6 mice were treated with NN and HH for 7 days, and the spleen was collected for subsequent detection. (A) Western blot detection of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression in spleen. (B) Statistical analysis of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression in A. (C) GEO data analysis of HMOX1, FTL, FTH1, NCOA4, SLC40A1, and TFRC mRNA expression in PMBCs before and after climbing to high altitude (n=98). (D) Western blot detection of ACSL4, GPX4 and xCT protein expression in the spleen. (E) Statistical analysis of ACSL4, GPX4 and xCT protein expression in D. (F) GEO data analysis of ACSL4, GPX4 and SLC7A11 mRNA expression in PMBCs before and after climbing to high altitude (n=98). (G) The content of Fe2+ in the spleen was detected using the FerroFarRed probe by flow cytometry. (H) The level of lipid ROS in the spleen was detected using the C11-BODIPY probe by flow cytometry. The MDA (I), Cys (J) and GSH (K) levels in the spleen were detected using biochemical detection kits. Data were expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: The C57BL/6 mice were treated with NN and HH for 7 days, and the spleen was collected for subsequent detection. (A) Western blot detection of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression in spleen. (B) Statistical analysis of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression. (C) GEO data analysis of HMOX1, FTL, FTH1, NCOA4, SLC40A1, and TFRC mRNA expression in PMBCs before and after climbing to high altitude (n=98). Western blot detection of ACSL4, GPX4, and xCT protein expression in the spleen is shown in (D) and its statistical analysis in (E). (F) GEO data analysis of ACSL4, GPX4 and SLC7A11 mRNA expression in PMBCs before and after reaching high altitude (n=98). (G) The content of Fe2+ in the spleen was detected using the FerroFarRed probe by flow cytometry. (H) The level of lipid ROS in the spleen was detected using the C11-BODIPY probe by flow cytometry. The MDA (I), Cys (J) and GSH (K) levels in the spleen were detected using biochemical detection kits, respectively. (L) Lillie staining of Fe2+ in the spleen after 7 days of HH exposure. (M) Fe2+ deposition in the spleen as described in (L) was quantified. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Added in the revised version: With the exception of NCOA4, the alterations in gene expressions related to iron metabolism and ferroptosis in PBMCs were consistent with our Western blot results (Figure 5, C and F). The GPX4 gene and protein expressions remained largely unchanged in both human PBMCs and mouse spleens. The changes in iron metabolism-and ferroptosis-related genes in PBMCs reflect the protein changes in the spleen under HH exposure. We detected Fe2+ and lipid ROS levels in the spleen by flow cytometry, and quantified MDA, GSH, and Cys levels using biochemical detection kits. As depicted in Figure 5, G and H; Figure S1, I-K, the Fe2+ (Figure 5G; Figure S1, E and F) and lipid ROS (Figure 5H; Figure S1, G and H) levels in the spleen increased significantly after HH exposure. Additionally, the MDA content (Figure 5I; Figure S1I) in the spleen increased, whereas Cys (Figure 5J; Figure S1J) and GSH (Figure 5K; Figure S1K) levels significantly decreased after HH exposure. To ascertain the increased Fe2+ primarily emanated from the red pulp, we detected Fe2+ deposition and distribution in the spleen by Lillie stain following 7 and 14 days of HH exposure (Figure 5, L-M; Figure S1, L-M). These results demonstrated increased Fe2+ primarily deposited in the red pulp of the spleen. Taken together, these findings suggest that iron mobilization in the spleen was enhanced and ferroptosis was induced after 7 days of HH exposure.

Earlier version: Hypoxia promotes ferroptosis of primary splenic macrophages

Revised version: Hypoxia induces ferroptosis in primary splenic macrophages

Earlier version: To study whether hypoxia exposure can induce ferroptosis in macrophages, changes in Fe2+ and lipid ROS levels, cell viability, phagocytosis and ferroptosis-related protein expression in primary splenic macrophages were detected under hypoxia in the presence of a ferroptosis inhibitor (Fer-1). As shown in Figure 6, compared with the normoxia control group (Nor), the levels of Fe2+ (Figure 6A-C) and lipid ROS (Figure 6D-E) were significantly increased, and the number of viable cells (Figure 6G-H) and phagocytosis ability (Figure 6I) were also significantly increased after 24 h of hypoxia (Hyp) treatment. In addition, with prolonged hypoxia time, the expression of ACSL4 gradually increased, while the expression of xCT and GPX4 decreased (Figure 6F). Fer-1 inhibited the increase in Fe2+ content (Figure 6J-K) and reversed the expression changes in ACSL4, xCT and GPX4 caused by hypoxia exposure (Figure 6L-O). At the same time, Fer-1 reversed the levels of MDA, Cys and GSH induced by hypoxia (Figure 6P-R). These results confirmed that hypoxia induced ferroptosis in primary splenic macrophages.

Revised version: To investigate whether hypoxia exposure can trigger ferroptosis in macrophages, we measured alterations in Fe2+ and lipid ROS levels, cell viability, phagocytosis, and ferroptosis-related protein expressions in primary splenic macrophages under hypoxia in the presence of a ferroptosis inhibitor (Fer-1). As depicted in Figure 6, compared with the normoxia control group (Nor), Fe2+ (Figure 6, A-C) and lipid ROS (Figure 6, D-E) levels significantly increased, whereas the number of viable cells (Figure 6G-H) and phagocytosis ability (Figure 6I) significantly decreased after 24 h of hypoxia (Hyp) treatment. In addition, prolonged hypoxia exposure led to a gradual increase in ACSL4 expression, while xCT and GPX4 expressions decreased (Figure 6F). Fer-1 mitigated the increase in Fe2+ content (Figure 6, J-K) and reversed the expression changes in ACSL4, xCT, and GPX4 induced by hypoxia exposure (Figure 6, L-O). Simultaneously, Fer-1 reversed the alterations in MDA, Cys, and GSH levels induced by hypoxia (Figure 6, P-R). These findings confirm that hypoxia induced ferroptosis in primary splenic macrophages.

Added in the revised version: Hypoxia enhances primary peritoneal macrophage ferroptosis.

Earlier version: Hypoxia promoted primary peritoneal macrophage ferroptosis. Primary peripheral macrophages were cultured under 1% hypoxia for different times (0, 12, 24, 48 h) with or without pre-treatment with 10 µM Fer-1 for 1 h, and then the cells were collected for subsequent detection. (A) Immunofluorescence detection of Fe2+ levels in macrophages after hypoxia exposure for 24 h using the FerroFarRed probe under a fluorescence microscope. (B) Quantitative results of fluorescence intensity in the macrophages described in A. (C) Flow cytometry detection of Fe2+ levels in macrophages exposed to hypoxia for 24 h using the FerroFarRed probe. (D) Immunofluorescence detection of lipid ROS levels in macrophages exposed to hypoxia for 24 h using the C11-BODIPY probe under a fluorescence microscope. (E) Quantitative results of mean fluorescence intensity in the macrophages described in D. (F) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages under hypoxia for different times. (G) Calcein/PI double staining in macrophages after 24 h of exposure to hypoxia was analysed by flow cytometry. (H) Cell death proportions in macrophages are given as bar graphs. (I) Macrophage phagocytic activity of Cy5.5-labelled Ecoil after 24 h of hypoxia exposure by flow cytometry in vitro. (J) Immunofluorescence detection of Fe2+ levels in macrophages pre-treated with Fer-1 followed by 24 h of hypoxia exposure using the FerroFarRed probe. (K) Quantitative results of fluorescence intensity in the macrophages described in J. (L) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages pre-treated with Fer-1 followed by 24 h of hypoxia exposure. (M and N) Statistical analysis of ACSL4, GPX4 and xCT protein expression in L. MDA (P), Cys (Q), and GSH (R) levels in macrophages pre-treated with Fer-1 followed by 24 h of hypoxia exposure were detected using kits by biochemical methods. Data were expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P <0.001 versus the NN group or the indicated group.

Revised version: Primary peripheral macrophages were cultured under 1% hypoxia for varying durations (0, 12, 24, 48 h), either with or without a pre-treatment of 10 µM Fer-1 for 1 h, before being collected for further examination. (A) Immunofluorescence detection of Fe2+ levels in macrophages after hypoxia exposure for 24 h using the FerroFarRed probe under a fluorescence microscope. (B) Quantitative results of fluorescence intensity in the macrophages described in A. (C) Flow cytometry detection of Fe2+ levels in macrophages exposed to hypoxia for 24 h using the FerroFarRed probe. (D) Immunofluorescence detection of lipid ROS levels in macrophages exposed to hypoxia for 24 h using the C11-BODIPY probe under a fluorescence microscope. (E) Quantitative results of mean fluorescence intensity in the macrophages described in D. (F) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages under hypoxia for different times. (G) Flow cytometry was employed to analyze Calcein/PI double staining in macrophages following a 24 h hypoxia exposure. (H) Cell death proportions in macrophages are given as bar graphs. (I) Macrophage phagocytic activity against Cy5.5-labelled E.coli following 24 h of hypoxia exposure was assessed using flow cytometry in vitro. (J) Fe2+ levels in macrophages pre-treated with Fer-1, then exposed to 24 h of hypoxia, were detected via the FerroFarRed probe in immunofluorescence. (K) Quantitative results of fluorescence intensity in the macrophages described in J. (L) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages pre-treated with Fer-1, then followed by 24 h of hypoxia exposure. (M-O) Statistical analysis of ACSL4, xCT, and GPX4 protein expression in L. MDA (P), Cys (Q), and GSH (R) levels in macrophages pre-treated with Fer-1 and subsequently exposed to 24 h of hypoxia were detected using kits and biochemical methods. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Discussion

The purpose of this study was to explore the effects and mechanisms of HA/HH exposure on erythrophagocytosis and iron circulation in mouse spleens. Here, we used an HH chamber to simulate 6000 m HA exposure and found that HH exposure induced iron mobilization and activated ferroptosis in the spleen, especially in red pulp macrophages (RPMs), which subsequently inhibited the phagocytosis and clearance of RBCs. Finally, chronic exposure to HA/HH may promote the retention of RBCs in the spleen, cause splenomegaly, advance RBC production, and promote the occurrence and development of HAPC (Figure 7).

Added in the revised version: High altitude/hypobaric hypoxia exposure suppresses erythrophagocytosis by inducing macrophages ferroptosis in the spleen.

Earlier version: High altitude/hypobaric hypoxia exposure inhibits erythrophagocytosis by inducing ferroptosis in macrophages in the spleen. HA/HH exposure induces iron mobilization, especially increased Fe2+ levels, which are further increased in the spleen/macrophages of mice. On the other hand, HA/HH exposure increased ACSL4 expression and inhibited GSH production by suppressing xCT expression and GSH production. The above results enhance lipid peroxidation and further macrophage ferroptosis. Meanwhile, HA/HH exposure decreases monocyte migration and differentiation from the BM to the spleen. Taken together, HA/HH exposure inhibits erythrophagocytosis and further promotes erythrocytosis, which may aggravate HAPC development.

Revised version: HA/HH exposure induces iron mobilization, especially increased Fe2+ levels, which are further increased in the spleen/macrophages of mice. On the other hand, HA/HH exposure increases ACSL4 expression, suppressing xCT expression and GSH production, which escalates lipid peroxidation and subsequent macrophage ferroptosis. Meanwhile, HA/HH exposure decreases monocyte migration and differentiation from the BM to the spleen. Taken together, HA/HH exposure inhibits erythrophagocytosis and further promotes erythrocytosis, which may aggravate HAPC development.

Earlier version: RBCs/erythrocytes are the carriers of oxygen and the most abundant cell type in the body. Erythrocytes are rapidly increased by triggering splenic contraction in acute HA/HH exposure [16] and stimulating erythropoiesis in subsequent continuous or chronic HA/HH exposure. Changes in spleen morphology and size are closely related to the spleen’s ability to recover RBCs [17], and studies have shown that the spleen is in a contraction state after short-term exposure to HA/HH, and approximately 40% of the increase in RBCs is due to the contraction of the spleen [16, 18, 19]. In the present study, mice were treated under HH to mimic HA exposure (mainly mimicking the low oxygen partial pressure environment of 6000 m HA) for different times according to other studies [20, 21]. Our results showed that HH exposure not only significantly increased the RBC and HGB contents but also significantly increased CD47+ cells in blood. However, short-term (1 d) HH exposure caused spleen contraction and induced splenomegaly after 3 d of HH treatment. These results are consistent with other research results; that is, HH or HA exposure affects spleen morphology and further affects RBC counts and HGB levels [22–24].

Revised version: RBCs/erythrocytes are the carriers of oxygen and the most abundant cell type in the body. Erythrocytes are rapidly increased by triggering splenic contraction in acute HA/HH exposure [21] and stimulating erythropoiesis in subsequent continuous or chronic HA/HH exposure. Changes in spleen morphology and size are closely related to the spleen’s ability to recover RBCs [22], and studies have shown that the spleen is in a contraction state after short-term exposure to HA/HH, and approximately 40% of the increase in RBCs is due to the contraction of the spleen [21, 23, 24]. In the present study, mice were treated under HH to mimic HA exposure (mainly mimicking the low oxygen partial pressure environment of 6000 m HA) for different times according to other studies [25, 26]. Our results showed that HH exposure not only significantly increased the RBC and HGB contents but also significantly increased CD47+ cells in blood. However, short-term (1 d) HH exposure caused spleen contraction and induced splenomegaly after 3 d of HH treatment. This contraction can trigger an immediate release of RBCs into the bloodstream in instances of substantial blood loss or significant reduction of RBCs. Moreover, elevated oxygen consumption rates in certain animal species can be partially attributed to splenic contractions, which augment hematocrit levels and the overall volume of circulating blood, thereby enhancing venous return and oxygen delivery [27, 28]. Thus, we hypothesized that the body, under such conditions, is incapable of generating sufficient RBCs promptly enough to facilitate enhanced oxygen delivery. Consequently, the spleen reacts by releasing its stored RBCs through splenic constriction, leading to a measurable reduction in spleen size. These results are consistent with other research results; that is, HH or HA exposure affects spleen morphology and further 483 affects RBC counts and HGB levels [29–31]. Considering these findings, the present study postulates that hypoxia-induced inhibition of erythrophagocytosis may lead to RBC retention. However, we acknowledge that the manuscript in its current preprint form does not offer conclusive evidence to substantiate this hypothesis. To bridge this gap, we further conducted experiments where the spleen was perfused, and total cells were collected post HH exposure. These cells were then smeared onto slides and subjected to Wright staining. Our results unequivocally demonstrate an evident increase in deformation and retention of RBCs in the spleen following 7 and 14 days of HH exposure. This finding strengthens our initial hypothesis and contributes a novel perspective to the understanding of splenic responses under hypoxic conditions.

Earlier version: HAPC is a common chronic HA disease characterized by excessive proliferation of RBCs caused by HH conditions [25]. Under physiological conditions, RBC homeostasis is maintained by balancing erythropoiesis and erythrocyte clearance [4]. The ability of RBC disposal in the spleen for iron recycling under HA/HH exposure was important to RBC regeneration in bone marrow (BM) [26]. Thus, we hypothesized that erythrophagocytosis and iron recycling in the spleen were altered by HA/HH exposure and further disturbed RBC homeostasis and affected the progression of HAPC. We found that compared with sham group mice, HH significantly increased the contents of RBCs and HGB in the blood of splenectomy mice. This strongly verified that the spleen is a key organism that maintains RBC magnitude within a certain range of physiological statuses under HA/HH conditions.

Revised version: HAPC is a common chronic HA disease characterized by excessive proliferation of RBCs caused by HH conditions [32]. Under physiological conditions, RBC homeostasis is maintained by balancing erythropoiesis and erythrocyte clearance [4]. The ability of RBC disposal in the spleen for iron recycling under HA/HH exposure was important to RBC regeneration in bone marrow (BM) [33]. Thus, we hypothesized that erythrophagocytosis and iron recycling in the spleen were altered by HA/HH exposure and further disturbed RBC homeostasis and affected the progression of HAPC. We found that compared with sham group mice, HH significantly increased the contents of RBCs and HGB in the blood of splenectomy mice. This strongly verified that the spleen is a key organism that maintains RBC magnitude within a certain range of physiological statuses under HA/HH conditions.

Earlier version: As originally proposed by Metchnikoff in the 19th century, macrophages, especially RPMs, play a pivotal role in the regulation of RBCs and iron homeostasis [27, 28]. We detected the macrophage population in the spleen and found that HH exposure for 7 and 14 days reduced the total macrophage and M1-and M2-type macrophage proportions in the spleen. This result is supported by another study, that is, hypoxia reduces M1-type macrophages in human gastric cancer [29]. We next observed whether monocyte migration and differentiation from the BM to the spleen replenish reduced macrophages and sustain macrophage homeostasis after HH treatment. It is well known that splenic erythropoiesis is a dynamic process [4]. Upon phagocytosing erythrocytes in macrophages, spleen tissue (including macrophages and fibroblasts) produces the chemokine C–C motif chemokine ligand 2 and 7 (CCL2 and CCL7) to recruit blood monocytes to the spleen [30, 31]. Unlike most studies in tumours, which mainly suggested that macrophage accumulation in hypoxic areas was derived by enhanced chemokine expression in tumours [32, 33], CCL2, CCL7 and Csf1 expression was decreased in the spleen after HH exposure for 7 and 14 days in our study. Furthermore, we found that HH exposure inhibited Ly6Chi monocyte migration from the BM to spleen, where these cells subsequently differentiated into RPMs. The combination of decreased splenic RPMs, along with the reduced BM-derived Ly6Chi monocytes, suggests that the homeostasis of the RPM population of circulating monocytes was also inhibited after the RPM depletion induced by HH exposure.

Revised version: As originally proposed by Metchnikoff in the 19th century, macrophages, especially RPMs, play a pivotal role in the regulation of RBCs and iron homeostasis [34, 35]. We detected the macrophage population in the spleen and found that HH exposure for 7 and 14 days reduced the total macrophage and M1-and M2-type macrophage proportions in the spleen. This result is supported by another study, that is, hypoxia reduces M1-type macrophages in human gastric cancer [36]. We next observed whether monocyte migration and differentiation from the BM to the spleen replenish reduced macrophages and sustain macrophage homeostasis after HH treatment. It is well known that splenic erythropoiesis is a dynamic process [4]. Upon phagocytosing erythrocytes in macrophages, spleen tissue (including macrophages and fibroblasts) produces the chemokine C–C motif chemokine ligand 2 and 7 (CCL2 and CCL7) to recruit blood monocytes to the spleen [37, 38]. Unlike most studies in tumours, which mainly suggested that macrophage accumulation in hypoxic areas was derived by enhanced chemokine expression in tumours [39, 40], CCL2, CCL7 and Csf1 expression was decreased in the spleen after HH exposure for 7 and 14 days in our study. Furthermore, we found that HH exposure inhibited Ly6Chi monocyte migration from the BM to spleen, where these cells subsequently differentiated into RPMs. The combination of decreased splenic RPMs, along with the reduced BM-derived Ly6Chi monocytes, suggests that the homeostasis of the RPM population of circulating monocytes was also inhibited after the RPM depletion induced by HH exposure. The observed decrease in monocytes within the BM is likely attributable to the fact that monocytes and precursor cells for RBCs both originate from the same hematopoietic stem cells within the BM [41]. As such, the differentiation to monocyte is reduced under hypoxic conditions, which may subsequently cause a decrease in migration to spleen. Furthermore, we hypothesize that an increased migration of monocytes to other tissues under HH exposure may also contribute to the decreased migration to the spleen. The liver, which partially contributes to the clearance of RBCs, may play a role in this process. Our investigations to date have indeed identified an increased monocyte migration to the liver. We were pleased to discover an elevation in CSF1 expression in the liver following HH exposure for both 7 and 14 days. This finding was corroborated through flow cytometry, which confirmed an increase in monocyte migration to the liver.

Earlier version: HO-1 is a major antioxidant and cytoprotective enzyme that catalyzes the degradation of heme to provide heme homeostasis and protect against free heme-induced toxicity [34]. In addition, it was supposed that mice deficient in HO-1 are largely devoid of RPMs and suggested that HO-1 is critical for RPM development and survival after erythrocyte clearance [35]. The reduced HO-1 protein of red pulp in our study also implied decreased macrophage accounts induced by HH exposure, although HO-1 was considered to be upregulated by hypoxia-inducible factor 1 (HIF-1) under hypoxia [36, 37]. Along with the decreased macrophage amount induced by HH, the phagocytic ability of macrophages in the spleen was also inhibited both in vitro and in vivo. Meanwhile, the injection of tufftsin strongly confirmed the enhancement of heme iron processing in the spleen under HH exposure. Nevertheless, our data conflict with the study of Anand et al, which reported that hypoxia causes an increase in phagocytosis by macrophages in a HIF-1α-dependent manner [38]. This may be because they performed intermittent hypoxia, while we used consistent hypoxia in the present study. The central role played by macrophages in supporting erythrocyte homeostasis stems, in part, from digesting the HGB content of clearance RBCs and recycling the iron back to erythroid progenitors for heme synthesis and HGB production [39]. In mammals, the majority of the body iron is in the form of heme, and the largest pool of heme consists of HGB [40]. HGB contains four prosthetic hemoglobins, and each mature RBC is thought to contain approximately 1.2 × 109 heme moieties [39]. As previously mentioned, the majority of iron required to sustain erythropoiesis is derived from recycled RBCs. In general, RPMs are equipped with molecular machinery capable of neutralizing the toxic effects of heme and metabolizing iron [41, 42]. Nevertheless, defects in erythrocytophagy in macrophages probably lead to aberrant iron metabolism, including anemia and iron overload [41]. The release of heme upon processed RBCs constitutes a permanent and considerable threat for iron cytotoxicity in macrophages and may eventually result in a specific form of programmed cell death termed ferroptosis [43], which may cause decreased cell counts [44, 45].

Revised version: HO-1 is a major antioxidant and cytoprotective enzyme that catalyzes the degradation of heme to provide heme homeostasis and protect against free heme-induced toxicity [42]. In addition, it was supposed that mice deficient in HO-1 are largely devoid of RPMs and suggested that HO-1 is critical for RPM development and survival after erythrocyte clearance [43]. The reduced HO-1 protein of red pulp in our study also implied decreased macrophage accounts induced by HH exposure, although HO-1 was considered to be upregulated by hypoxia-inducible factor 1 (HIF-1) under hypoxia [44, 45]. Along with the decreased macrophage amount induced by HH, the phagocytic ability of macrophages in the spleen was also inhibited both in vitro and in vivo. Meanwhile, the injection of tufftsin strongly confirmed the enhancement of heme iron processing in the spleen under HH exposure. Nevertheless, our data conflict with the study of Anand et al, which reported that hypoxia causes an increase in phagocytosis by macrophages in a HIF-1α-dependent manner [46]. This may be because they performed intermittent hypoxia, while we used consistent hypoxia in the present study. Macrophages play a crucial role in maintaining erythrocyte homeostasis, partially due to their function in digesting the HGB content of cleared RBCs and recycling the iron back to erythroid progenitors for heme synthesis and HGB production [47]. In mammals, the majority of bodily iron exists in the form of heme, with the most substantial pool comprising HGB [48]. HGB contains four prosthetic hemoglobins, and each mature RBC is thought to contain approximately 1.2 × 109 heme moieties [47]. As previously mentioned, the majority of iron required to sustain erythropoiesis is derived from recycled RBCs. In general, RPMs are equipped with molecular machinery capable of neutralizing the toxic effects of heme and metabolizing iron [49, 50]. Nonetheless, any disruptions in the process of erythrophagocytosis, the uptake and degradation of erythrocytes by macrophages, could potentially result in abnormal iron metabolism, potentially manifesting as conditions such as anemia or iron overload [49]. The release of heme upon processed RBCs constitutes a permanent and considerable threat for iron cytotoxicity in macrophages and may eventually result in a specific form of programmed cell death termed ferroptosis [51], which may cause decreased cell counts [52, 53].

Earlier version: Accordingly, we investigated the iron metabolism status and explored ferroptosis in spleen/macrophages after HH/hypoxia treatment and found that HH exposure prompted iron mobilization, especially in ferritinophagy and lipid peroxidation in the spleen. Interestingly, we found that except for NCOA4, all the gene expression changes in the PBMCs of humans after acute HA exposure were similar to the changes in the spleens of mice after HH exposure for 7 and 14 days. This is probably because NCOA4 in the spleen mediates iron release from Ft-L and facilitates iron reuse, while PBMCs do not possess this function in vivo. These results not only indicated that HH exposure induces spleen ferroptosis but also implied that the gene expression changes in PBMCs under HA/HH may reflect RBC processing functions in the spleen and further indicate the iron metabolism status in the clinic. We further found that the exact mechanism of macrophage ferroptosis induced by HH exposure was caused by increased Fe2+ and decreased antioxidative system expression, which finally resulted in lipid peroxidation of macrophages. It has been proposed that heme catabolism by HO-1 protects macrophages from oxidative stress [41]. The enhanced ROS and lipid peroxidation in vivo were also consistent with the decreased HO-1 expression in the spleen. In addition, 1% hypoxia treatment induced ferroptosis in vitro, which was reduced by treatment with ferrostatin-1, a ferroptosis inhibitor. However, our data were inconsistent with the study of Fuhrmann et al., which reported that hypoxia inhibits ferritinophagy and protects against ferroptosis [46]. We hypothesized that the enhanced iron demand for new RBC regeneration in BM caused by HH leads to an increase in NCOA4 and Ft-L protein expression in the spleen. On the other hand, hypoxia inhibited anti-ferroptosis system expression, including reduced Gpx4 expression and increased lipid ROS production. The results were supported by the group of Youssef LA et al, who reported that increased erythrophagocytosis induces ferroptosis in RPMs in a mouse model of transfusion [47]. However, as most studies have reported, ferroptosis is not only characterized by increased lipid ROS but also significantly shrinking mitochondria [48]. We never found shrinking mitochondria in RPMs after HH exposure (data not shown). Gao et al. proposed that mitochondria played a crucial role in cysteine deprivation-induced ferroptosis but not in that induced by inhibiting glutathione peroxidase-4 (GPX4), the most downstream component of the ferroptosis pathway [49]. Interestingly, in our in vivo study, Gpx4 expression was also not changed after HH exposure. It is still not clear whether mitochondria are involved in ferroptosis. Whether HH treatment caused mitochondrial swelling directly and the exact mechanism involved in this process still need to be further investigated.

Revised version: Accordingly, we investigated the iron metabolism status and explored ferroptosis in spleen/macrophages after HH/hypoxia treatment and found that HH exposure prompted iron mobilization, especially in ferritinophagy and lipid peroxidation in the spleen. Interestingly, we found that except for NCOA4, all the gene expression changes in the PBMCs of humans after acute HA exposure were similar to the changes in the spleens of mice after HH exposure for 7 and 14 days. This is probably because NCOA4 in the spleen mediates iron release from Ft-L and facilitates iron reuse, while PBMCs do not possess this function in vivo. These results not only indicated that HH exposure induces spleen ferroptosis but also implied that the gene expression changes in PBMCs under HA/HH may reflect RBC processing functions in the spleen and further indicate the iron metabolism status in the clinic. We further found that the exact mechanism of macrophage ferroptosis induced by HH exposure was caused by increased Fe2+ and decreased antioxidative system expression, which finally resulted in lipid peroxidation of macrophages. It has been proposed that heme catabolism by HO-1 protects macrophages from oxidative stress [49]. The enhanced ROS and lipid peroxidation in vivo were also consistent with the decreased HO-1 expression in the spleen. In addition, 1% hypoxia treatment induced ferroptosis in vitro, which was reduced by treatment with ferrostatin-1, a ferroptosis inhibitor. However, our data were inconsistent with the study of Fuhrmann et al., which reported that hypoxia inhibits ferritinophagy and protects against ferroptosis [54]. We hypothesized that the enhanced iron demand for new RBC regeneration in BM caused by HH leads to an increase in NCOA4 and Ft-L protein expression in the spleen. On the other hand, hypoxia inhibited anti-ferroptosis system expression, including reduced Gpx4 expression and increased lipid ROS production. The results were supported by the group of Youssef LA et al, who reported that increased erythrophagocytosis induces ferroptosis in RPMs in a mouse model of transfusion [55]. However, as most studies have reported, ferroptosis is not only characterized by increased lipid ROS but also significantly shrinking mitochondria [56]. We never found shrinking mitochondria in RPMs after HH exposure (data not shown). Gao et al. proposed that mitochondria played a crucial role in cysteine deprivation-induced ferroptosis but not in that induced by inhibiting glutathione peroxidase-4 (GPX4), the most downstream component of the ferroptosis pathway [57]. Interestingly, in our in vivo study, Gpx4 expression was also not changed after HH exposure. It is still not clear whether mitochondria are involved in ferroptosis. Whether HH treatment caused mitochondrial swelling directly and the exact mechanism involved in this process still need to be further investigated.

Added in the revised version: When the splenic macrophages’ phagocytic function is compromised, the burden of managing the increased number of RBCs falls to the other cells within the spleen, as well as other clearance mechanisms within the body [58]. The endothelium, for instance, has been shown to be involved in the removal of aged and damaged RBCs [59]. The sinusoidal endothelial cells can recognize these RBCs, which often have altered surface properties, leading to their removal from circulation [59, 60]. Therefore, under HH conditions, endothelial cells may play a compensatory role in RBC clearance when macrophage function is impaired. Other macrophages, such as liver Kupffer cells and bone marrow macrophages, may also contribute to managing this buildup of RBCs [61]. Kupffer cells are the resident macrophages of the liver and have significant phagocytic ability. Similarly, macrophages in the bone marrow play a crucial role in iron recycling from aged erythrocytes [61, 62]. Therefore, these macrophages could potentially compensate for the reduction in splenic erythrophagocytosis, although their ability to effectively handle a large surge in RBC numbers remains to be determined.

Added in the revised version: It is also possible that the spleen may adapt to the increased load of RBCs by expanding the red pulp, leading to splenomegaly [63]. The red pulp is highly vascular and contains numerous sinuses lined with macrophages [64]. This expansion could potentially accommodate a greater number of RBCs, thus enhancing the spleen’s capacity for erythrophagocytosis and iron recycling. However, these compensatory mechanisms may be insufficient or become overwhelmed over time, leading to pathological conditions such as HAPC. Further research is necessary to fully understand the consequences of splenic ferroptosis on RBC homeostasis and how these effects could be mitigated or reversed. Finally, as these observations are based on experimental models, they should be corroborated in human studies to assess their clinical relevance and potential implications for managing conditions related to high altitude exposure. The potential for therapeutic interventions that target ferroptosis, erythrophagocytosis, and iron recycling should also be explored, as these could provide new avenues for treating HAPC and other conditions related to high altitude exposure.

Earlier version: Based on the above experiments, we first put forward the “spleen theory” of HAPC (Figure 7). It hypothesizes that HH exposure induced spleen ferroptosis, especially RPMs, which further inhibited erythrophagocytosis, RBC clearance, HGB processing and iron recycling. Therefore, chronic exposure to HA/HH promotes the retention of RBCs in the spleen, leading to splenomegaly and the continuous production of RBCs. These findings may be clinically relevant to the pathological conditions of HAPC progression.

Revised version: Based on the above experiments, we first put forward the “spleen theory” of HAPC (Figure 7). It hypothesizes that HH exposure induced spleen ferroptosis, especially RPMs, which further inhibited erythrophagocytosis, RBC clearance, HGB processing and iron recycling. This theory identifies 637 a novel pathophysiological process in which HH exposure leads to the induction of spleen ferroptosis, particularly in red pulp. As a result, erythrophagocytosis, RBCs clearance, and iron recycling are impeded, causing a buildup of RBCs within the spleen, leading to splenomegaly and continued RBC production. The “spleen theory” of HAPC provides an intriguing new perspective on the physiological responses to hypoxia and highlights the role of the spleen as a vital organ in managing erythrocyte homeostasis under high altitude/hypoxic conditions. These findings may be clinically relevant to the pathological conditions of HAPC progression.

Abbreviations

BM: bone marrow; Cys: Cysteine; GSH: Glutathione; Ft-L: Ferritin light chains; Fpn: Ferroportin; Ft: Ferritin; HH: Hypobaric hypoxia; HA: High-altitude; HAPC: High-altitude polycythemia; HO-1: Heme oxygenase-1; HCT: Haematocrit; HGB: Haemoglobin; HIF-1a: Hypoxia inducible factor 1 a; MCH: Mean corpuscular haemoglobin; MDA: Malondialdehyde; NN: Normobaric normoxia; NCOA4: nuclear receptor coactivator 4; PBMCs: peripheral blood mononuclear cells; RBC: Red blood cell; RPM: Red pulp macrophage; TfR1: Transferrin receptor 1; Tf: Transferrin; TfR: Transferrin receptor.

Author contributions

Earlier version: QQL and GHW conceived, organized, and designed the study. YQG supervised the work. WPY, MQL, JD and JL performed the experiments. WPY, GW and BL contributed to the analysis of data. QQL and GHW prepared, wrote and revised the manuscript. All authors contributed to read, and approved the submitted version.

Revised version: QQL and GHW conceived, organized, and designed the study. YQG supervised the work. WPY, MQL, JD and JYL performed the experiments. WPY, GW and BL contributed to the analysis of data. QQL and GHW prepared, wrote and revised the manuscript. All authors contributed to read, and approved the submitted version.

Funding

Earlier version: This work was supported by Natural Science Foundation of China (Grants 32271228, 82171190, and 81873924), Open Cooperation Program from Key Laboratory of Extreme Environmental Medicine, Ministry of Education (KL2019GY011), China Postdoctoral Science Foundation (2020M673649), High-level Innovation and Entrepreneurship Talents Introduction Program of Jiangsu Province of China, and Nantong Municipal Science and Technology Project (MS12021020 and MS22021010).

Revised version: This work was supported by Natural Science Foundation of China (Grants 32271228, 82171190, and 81873924), Open Cooperation Program from Key Laboratory of Extreme Environmental Medicine, Ministry of Education (KL2019GY011), China Postdoctoral Science Foundation (2020M673649), Cultivate Candidate of the Jiangsu Province “333” Project, and Nantong Municipal Science and Technology Project (MS12021020 and MS22021010).

Availability of data and materials

The datasets used and/or analyzed during the current study are available from the corresponding author on reasonable request.

Declarations

Ethics approval and consent to participate

All animal care and experimental protocols were carried out according to the Chinese Animal Management Rules of the Ministry of Health and were authorized by the Animal Ethics Committees of Nantong University research program protocol #S20190219-011. The results of the study are presented clearly, honestly, and without fabrication, falsification, or inappropriate data manipulation.

Consent for publication

Not applicable.

Competing interests

The authors have no relevant financial or non-financial interests to disclose.

Author details

1Department of Physiology and Hypoxic Biomedicine, Institute of Special Environmental Medicine and Co-innovation Center of Neuroregeneration, Nantong University, 9 Seyuan Road, Chongchuan District, Nantong, Jiangsu 226019, China. 2College of High-Altitude Military Medicine, Institute of Medicine and Hygienic Equipment for High Altitude Region, Army Medical University, Chongqing 400038, China. 3Key Laboratory of Extreme Environmental Medicine and High-Altitude Medicine, Ministry of Education of China, Chongqing 400038, China. 4Department of Neurosurgery, Southwest Hospital, Army Medical University, Chongqing, Chongqing 400038, China.

Added in the revised version: Supporting information